Effects of nicotine compared to placebo gum on sensitivity to pain and mediating effects of peak alpha frequency

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public Review):

Summary:

In this study, Millard and colleagues investigated if the analgesic effect of nicotine on pain sensitivity, assessed with two pain models, is mediated by Peak Alpha Frequency (PAF) recorded with resting state EEG. The authors found indeed that nicotine (4 mg, gum) reduced pain ratings during phasic heat pain but not cuff pressor algometry compared to placebo conditions. Nicotine also increased PAF (globally). However, mediation analysis revealed that the reduction in pain ratings elicited by the phasic heat pain after taking nicotine was not mediated by the changes in PAF. Also, the authors only partially replicated the correlation between PAF and pain sensitivity at baseline (before nicotine treatment). At the group-level no correlation was found, but an exploratory analysis showed that the negative correlation (lower PAF, higher pain sensitivity) was present in males but not in females. The authors discuss the lack of correlation.

In general, the study is rigorous, methodology is sound and the paper is well-written. Results are compelling and sufficiently discussed.

Strengths:

Strengths of this study are the pre-registration, proper sample size calculation, and data analysis. But also the presence of the analgesic effect of nicotine and the change in PAF.

Weaknesses:

It would even be more convincing if they had manipulated PAF directly.

We thank Reviewer #1 for their positive and constructive comments regarding our study. We appreciate the view that the study was rigorous and methodologically sound, that the paper was well-written, and that the strengths included our pre-registration, sample size calculation, and data analysis.

In response to the reviewer's comment about more directly manipulating Peak Alpha Frequency (PAF), we agree that such an approach could provide a more direct investigation of the role of PAF in pain processing. We chose nicotine to modulate PAF as the literature suggested it was associated with a reliable increase in PAF speed. As mentioned in our Discussion, there are several alternative methods to manipulate PAF, such as non-invasive brain stimulation techniques (NIBS) like transcranial alternating current stimulation (tACS) or neurofeedback training. These approaches could help clarify whether a causal relationship exists between PAF and pain sensitivity. Although methods such as NIBS still require further investigation as there is little evidence for these approaches changing PAF (Millard et al., 2024).

Reviewer #2 (Public Review):

Summary:

The study by Millard et al. investigates the effect of nicotine on alpha peak frequency and pain in a very elaborate experimental design. According to the statistical analysis, the authors found a factor-corrected significant effect for prolonged heat pain but not for alpha peak frequency in response to the nicotine treatment.

Strengths:

I very much like the study design and that the authors followed their research line by aiming to provide a complete picture of the pain-related cortical impact of alpha peak frequency. This is very important work, even in the absence of any statistical significance. I also appreciate the preregistration of the study and the well-written and balanced introduction. However, it is important to give access to the preregistration beforehand.

Comment on latest version:

I thank the authors for their efforts to respond to my previous concerns. I appreciate the additional analyses and the more cautious wording in several places. However, I still think the manuscript would benefit from a more direct interpretation of the findings.

Reviewer #3 (Public Review):

In this manuscript, Millard et al. investigate the effects of nicotine on pain sensitivity and peak alpha frequency (PAF) in resting state EEG. To this end, they ran a pre-registered, randomized, double-blind, placebo-controlled experiment involving 62 healthy adults who received either 4 mg nicotine gum (n=29) or placebo (n=33). Prolonged heat and pressure were used as pain models. Resting state EEG and pain intensity (assessed with a visual analog scale) were measured before and after the intervention. Additionally, several covariates (sex at birth, depression and anxiety symptoms, stress, sleep quality, among others) were recorded. Data was analyzed using ANCOVAequivalent two-wave latent change score models, as well as repeated measures analysis of variance. Results do not show ‘experimentally relevant’ changes of PAF or pain intensity scores for either of the prolonged pain models due to nicotine intake.

The main strengths of the manuscript are its solid conceptual framework and the thorough experimental design. The researchers make a good case in the introduction and discussion for the need to further investigate the association of PAF and pain sensitivity. Furthermore, they proceed to carefully describe every aspect of the experiment in great detail, which is excellent for reproducibility purposes. Finally, they analyse the data from almost every possible angle and provide an extensive report of their results.

The main weakness of the manuscript is the interpretation of these results. Even though some of the differences are statistically significant (e.g., global PAF, pain intensity ratings during heat pain), these differences are far from being experimentally or clinically relevant. The effect sizes observed are not sufficiently large to consider that pain sensitivity was modulated by the nicotine intake, which puts into question all the answers to the research questions posed in the study.

We would like to express our gratitude to Reviewer #3 for their thoughtful and constructive review, including the positive feedback on the strengths of our study's conceptual framework, experimental design, and thorough methodological descriptions.

We acknowledge the concern regarding the experimental and clinical relevance of some statistically significant results (e.g., global PAF and pain intensity during heat pain) and agree that small effect sizes may limit their practical implications. However, our primary goal was to assess whether nicotine-induced changes in PAF mediate pain changes, rather than to demonstrate large direct effects on pain sensitivity. Nicotine was chosen for its known ability to modulate PAF, and our focus was on the mechanistic role of PAF in pain perception. To clarify this, we have revised the discussion to better differentiate between statistical significance, experimental relevance, and clinical applicability. We emphasize that this study represents a preliminary step towards understanding PAF’s mechanistic role in pain, rather than a direct clinical application.

We appreciate the suggestion to refine our interpretation. We have adjusted our language to ensure it aligns with the effect sizes observed and made recommendations for future research, such as testing different nicotine doses, to potentially uncover stronger or more clinically relevant effects.

Although modest, we believe these findings offer valuable insights into the potential mechanisms by which nicotine affects alpha oscillations and pain. We have also discussed how these small effects could become more pronounced in different populations (e.g., chronic pain patients) and over time, offering guidance for future research on PAF modulation and pain sensitivity.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

I thank the authors for their efforts to respond to my concerns. I very much admire the thoughts and effort that the authors have put into this study. This is a great study design, a fantastic research idea and an important path to study the neurophysiology of pain.

I also understand that some of my suggestions are not compatible with the authors' current line of analysing their data. Since the authors did not change the analysis procedure and obviously could not change their experimental design, I would largely keep my public review.

However, acknowledge the different opinions and have a few remarks and suggestions.

(1) In agreement with reviewer 3, I would not consider the effect of nicotine to be meaningful. The random effect mentioned by the authors could have been addressed with a crossover design. In this case, the placebo and intervention could have been directly compared. Interindividual differences in pain perception are at least partially controlled for by analysing the difference between pre and post-gum. The distribution of the pre-post difference separately for placebo and nicotine would be a valuable addition to the manuscript. I still think the minimal effect of nicotine on pain is not meaningful and not interpretable. The relevance of any explanatory variable like stress could be by chance (PMID: 35296861). Reviewer 3 had a valid point (comment 6) and this is something that should be mentioned as clearly as possible: there was no effect of nicotine. For PHP the comparison with the placebo is irrelevant if there is no pre vs. post-effect.

It is fair to say that some issues could have been resolved using a cross-over design, however, the choice of a parallel design was made, in part, due to the desire to mask the presence of a placebo gum. Having a single session per participant allowed us to produce a deception that all participants were receiving nicotine gum, which was later revealed during a debrief. This aspect is far less feasible in a cross-over design. In future, different approaches could be used. And we appreciated the reviewers recognition that this cannot, of course, be changed at this stage.

Regarding the comments of 1) ’I would not consider the effect of nicotine to be meaningful’, 2) the minimal effect of nicotine on pain is not meaningful and not interpretable’, and 3) ’there was no effect of nicotine’, we discern that the latter differs considerably from the first two points. We would like to address this point by first calling into question what the reviewer(s) mean by meaningful, and highlighting that discerning this varies by context. Edenberg (2021) highlights that meaning can be situated within several contexts: statistical, biological (mechanistic), clinical, and public policy. We show the statistical context, and therefore would like to emphasise that we cannot state number 3) ’there was no effect of nicotine’.

Although the complexity of our analysis is not at the same scale, we appreciate the need for thousands of participants for brain-wide associations, as cited (PMID: 35296861 = Marek et al., 2022). The article referenced states that reproducible brain-wide association studies require thousands of individuals, and thus we would like to highlight that our work indicating the potential relevance of sex/stress greatly differs from this. Though it may be by chance, the potential for that chance is 0.05 from a statistical significance perspective. Certainly, no finding should be taken as truth without replication, but that does not mean we should not report our finding of a statistically significant effect of stress on change in PAF in our model.

Then, the question becomes, are we already at the point where the issue of ’how to go beyond statistical significance’ should be addressed? One could argue no, as this is the first study to assess this question of nicotine’s effect on PAF and pain at the same time. But we would like to argue yes, although only to the mechanistic/biological level, because the question of what is meaningful is dependent on the study goals (Edenberg, 2021). We aimed to gain mechanistic insight into the PAF-pain relationship.

We hoped to achieve some biological/mechanistic meaning in the present study, however, as outlined in the discussion, with small effect sizes for nicotine on PAF and nicotine on heat pain, and lack of a mediating effect, we may not have achieved this. Although, if replicated, even small magnitude effects can shape our understanding of relevant mechanisms underlying pain and aid future clinical perspectives (Edenberg, 2021; Ross & Bassett, 2024). Meaningfulness in clinical or public policy contexts should of course be in mind, but were clearly not the focus of this present work.

We have edited the manuscript to demonstrate the fact that our effect size is small.

We corrected the rounding to 2 decimals, and now also refer the reader to the supplementary materials table:

“Change in pain was reduced by a factor of -0.68 for the nicotine group compared to the placebo group when controlling for confounding variables (Figure 3). See Supplementary Material 3.2, Table 5, for full model output.”

We edited the abstract and the first paragraph of the discussion to make clear to the reader that the effect sizes were small. In our desire to keep the text clear and readable, we did not intend to overstretch our claims, and apologise if the manuscript was unintentionally portraying this overstretch. These are very important sections of the manuscript in which to mention such nuances, as these are the parts viewed most when readers do not wish to read the entire manuscript.

Abstract

“The nicotine group showed a small decrease in heat-pain ratings compared to placebo group when controlling confounders, and a small increase in PAF across the scalp from pre- to post-gum, both with and without confounder adjustment. These effects were most pronounced in the central-parietal and right-frontal electrodes.”

Discussion first paragraph

“Compared to placebo, our data demonstrate three key findings: 1) nicotine increased PAF speed; 2) nicotine reduced prolonged heat pain intensity but only when controlling for confounding and not for the prolonged pressure pain model, and 3) decreases in prolonged heat pain intensity were not mediated by changes in PAF. This suggests that changes in heat pain due to nicotine occur through mechanisms unrelated to change in PAF and require replication as the effects required control for confounding and the reduction was small in magnitude.”

Discussion section on effect of nicotine on PAF:

“In line with previous literature in non-smokers [38,47,48], we contribute additional evidence that chewing 4 mg of nicotine gum produces a statistically significant increase in PAF speed in non-smoking, nicotine-gum-naïve, pain-free, adult participants. An exploratory cluster-based permutation analysis indicated that the change in wide-window PAF (8–12 Hz) produced by nicotine, was most pronounced in a cluster of electrodes in the central-parietal region and a separate cluster of electrodes in the right-frontal region, rather than in the pre-registered sensorimotor ROI (i.e., Cz, C1, C2, C3, C4).”

Discussion section on the effect of nicotine on pain:

“While acknowledging the modest effect size, it’s essential to consider the context of our study’s focus. Assessing the clinical relevance of pain reduction is pertinent in applications involving the use of any intervention for pain management [71]. However, from a mechanistic standpoint, particularly in understanding the implications of and relation to PAF, the specific magnitude of the pain effect becomes less pivotal, as small effects can still have mechanistic or biological meaning without establishing clinical meaning [72]. Nevertheless, future research should examine whether effects on pain increase in magnitude with different nicotine administration regimens (i.e. dose and frequency).”

Discussion section on relationship between PAF and pain:

“However, as this is the first study attempting to directly manipulate PAF, observing only small effects on pain, we cannot confidently state that a causal relationship between PAF and pain sensitivity does not exist, and future research should explore these relationships with larger sample sizes, larger doses of nicotine, and alternative PAF modulators.”

Discussion final paragraph:

“We have shown that, compared to placebo, nicotine gum produced statistically significant increases global PAF speed and decreases pain ratings during prolonged heat pain. However, the effects of nicotine on heat pain were not mediated by changes in PAF. Therefore, we find no evidence that PAF is causally related to pain. The implication of this is that changing PAF may not be a useful target for pain modulation, because pain sensitivity may not be dependent on state changes in PAF. However, this is the first study to test modulation of PAF as a mechanism for reducing pain intensity. Moreover, while effects on PAF appear small but robust, effects on heat pain required control for confounding to be observed. More randomised experiments on the effects of nicotine that use larger sample sizes, higher nicotine doses, and pharmacokinetic tracking, whilst also reducing mediator measurement error are needed to confirm the present findings and uncover whether these factors account for the lack of mediation effect seen here.”

(2) I still think the authors made a mistake in using nicotine. Their new wording did not convince me, and I doubt that any reader would buy the "mediation by design" phrase.

We addressed the reviewers previous comment (comment 1) on nicotine being a ‘mistake’, and no longer suggest there is a “large body of evidence suggests that nicotine is an ideal choice for manipulating PAF”. We have been clearer about our intention to assess mediation-by-design, with the manipulation of PAF being out primary aim. For example, this edit in the discussion:

“The choice of nicotine was driven by its potential to change PAF, but also, short-term nicotine use is thought to have acute analgesic properties in experimental settings, with a review reporting that nicotine increased pain thresholds and pain tolerance [51].”

In addition, we discussed the ’direct’ manipulation of PAF using non-invasive brain stimulation methods, and have now added further future research directions using methods other than nicotine or NIBS:

“Furthermore, there may have been other mediators that suppress the mediating effects of PAF within our model [66,82,83]. This is plausible, as nicotine is thought to act on pain through multiple pathways [52,54,84], and more direct PAF modulation could more definitively evaluate whether PAF is mechanistically linked to pain sensitivity. Future interventions to modulate PAF could use non-invasive brain stimulation (NIBS), such as transcranial alternating current stimulation (tACS) or repetitive transcranial magnetic stimulation (rTMS), to influence neuronal firing and thereby directly target the brain mechanisms theorised to generate PAF [85–87]. This approach has since been investigated following the initial preprint of

the present manuscript [88,89]. In addition, other approaches to modulating approaches could be considered, such as exercise [32–34], visual stimulation [35,36], cannabis [90,91], and neurofeedback [92,93], while considering their viability in terms of blinding, tolerability, and scalability.”

(3) I acknowledge the work of the authors on SNR, the frequency spectrum looks smooth enough to detect peaks. However, the problem of different alpha sources is not solved. This can be seen in the frequency spectrum, where some participants show two peaks. This is not the first time I have suggested this to the authors; the idea is not new. The fMRI analogue would be to average all GM voxels; instead of global alpha it would be "global brain activity". Some suggestions for a future approach that the authors might consider:

Run a group ICA and select topographies from there. This should simplify the (potentially arbitrary) selection process.

Consider filtering the data before running the ICA, at least >20 Hz to exclude muscle artefacts.

We respectfully note that the reviewer’s comment appears to overlook the substantial revisions we have already made in direct response to this issue. Specifically, we implemented ICA in the revised manuscript, using an automated, objective pipeline to identify components with a clear alpha peak and a topography aligned with a predefined sensorimotor template. This was explicitly done to address concerns about overlapping alpha sources and improve anatomical specificity.

Contrary to the reviewer’s implication, we did not rely solely on global analyses in the original or revised manuscript, as we also included pre-defined ROIs and data-driven ROIs from cluster-based permutation analysis. The additional ICA-based analyses—included in the first revision—produced results that were consistent with the original sensor-space ROI findings, supporting the robustness of the PAF–pain mediation effect regardless of the method used to isolate sensorimotor alpha.

We agree that group ICA could be a valuable complementary approach. However, it is not without limitations, particularly in EEG where inter-individual variability in source topographies is high. Our choice to perform ICA at the individual level and apply an automated selection criterion based on spectral and spatial features was deliberate, data-driven, and avoids the subjectivity involved in manual selection or group-level back-projection.

We have added a discussion of this point and fully acknowledge that future work may benefit from comparing multiple alpha sources (e.g., mu vs. occipital) more explicitly. However, we must emphasise that this is not a missing analysis—it is already included in the revised manuscript. In this second revision, we have added an additional reference to this supplementary analysis within the section previously titled “Exploratory cluster-based permutation anlaysis”, now titled “Exploratory data-driven analysis of alpha”, in case the previous reference to it was in a sub-optimal location:

“Lastly, use of in-house code developed to automatically extract a sensorimotor alpha component using independent component analysis (ICA) showed similar results to the sensorimotor ROI mediation analysis (Supplementary Material 3.11). This supports the robustness of the PAF-pain effect regardless of the method used to isolate sensorimotor alpha.”

We strongly disagree with the suggestion to apply a low-pass filter below 20 Hz prior to ICA. This would remove high-frequency content, including muscle activity—the very signals ICA is designed to isolate and separate from neural activity.

By removing muscle-related frequencies before ICA, the reviewer’s proposed approach would prevent ICA from identifying muscle components in the data. This is methodologically flawed. as ICA works best when it has access to the full spectral range of the signal, allowing it to separate statistically independent sources— including muscle, eye movements, and neural oscillations—based on their spatial and temporal structure (Onton et al., 2006). Thus the desired sensorimotor alpha component would be more distinct if the full spectral range is included. Low-pass filtering the data before ICA designed to identify alpha would severely compromise artifact removal, degrade component separation, and reduce the effectiveness of our preprocessing pipeline.

References for responses to reviewer 2

Edenberg HJ. Perspective on beyond statistical significance: Finding meaningful effects. Complex Psychiatry 2021;7:1–8. https://doi.org/10.1159/000517237.

Marek S, Tervo-Clemmens B, Calabro FJ, Montez DF, Kay BP, Hatoum AS, Donohue MR, Foran W, Miller RL, Hendrickson TJ, Malone SM. Reproducible brain-wide association studies require thousands of individuals. Nature. 2022 Mar 24;603(7902):654-60.

Ross LN, Bassett DS. Causation in neuroscience: keeping mechanism meaningful. Nature Reviews Neuroscience. 2024 Feb;25(2):81-90.

Reviewer #3 (Recommendations for the authors):

I would like to thank the authors for the time and dedication that they clearly put into addressing the reviewer's comments. As researchers, we alternate between both sides (as authors and reviewers) and we are all privy of how difficult these tasks can be.

For me as a reviewer, this is one of the tough times. I wanted to support this manuscript because it is clear that conducting the study had to be a monumental effort, and there were early indications that it was rigorously carried out. Specifically referring to rigor, I value preregistration as a tool to assess the strength of the hypothesis posed before conducting the study, and the severity with which the claims are tested. During my first review, I tried to access the preregistration and even provided my email address as a sign of transparency, but access was not granted. During my second review, the main file in the submission system contained two versions of the same manuscript (the first one corresponding to the old version with the wrong preregistration link). I was not initially aware of this and I apologize for the misunderstanding.

With respect to the rest of my observations, I see that the authors made concessions on small and less relevant issues (e.g. comments #1, #2, #3), but did not satisfactorily address the main ones (#4, #5, #6, #7, #8 and #9). Addressing those comments (which are in line with the issues noted by other reviewers), would probably have lead to a different interpretation of the results compared with what they state in the first paragraph of the discussion.

The main claims (abstract and first paragraph of the discussion) are " 1) nicotine increased PAF speed; 2) nicotine reduced prolonged heat pain intensity but not prolonged pressure pain intensity compared to placebo gum". Just by reading and contrasting the claims with the results, my assessment is that the strength of evidence is inadequate.

I was now able to review the analysis plan, which did not contain an estimation of the effect sizes (only effect directions) and did not have a proper sample size estimation (comment #4) and the main outcome was not, to my understanding, based in terms of the beta coefficients (as you replied to reviewer #1). I also do not think that the possibility of no effect or a false positive (comment #5) is fairly weighed in relation to these main claims in the abstract an first paragraph of the discussion.

I also do not consider the response to comment #6 as satisfactory, first because it is a mistake to consider that we (as reviewers) only interpret the 'mean results'. As we have detailed access to the data, we interpret these results in terms of signal-to-noise ratio, i.e., how large is the effect (signal) compared with the intra and interindividual variance (noise), and it should be clear for anyone working on this fields (pain and EEG) that not only the main effect is not experimentally relevant to support the claims (comments #6 and #7), but in some cases it is likely to have arisen from the unjustified inclusion of covariates (comment #8 and #9). I only raise the final strength of evidence in my assessment from 'inadequate' to 'incomplete', because of the additions to the discussion.

As a final thought, I would have been more than happy to support and recommend the reading of a manuscript that described a carefully designed and preregistered study, but in which the hypothesis were not corroborated by the data. I would have even marked the strength of evidence as compelling. Unfortunately, the preregistered hypothesis cannot be severely tested (at least not in terms of a frequentist framework) because the effects sizes were not predefined and the required sample size was not estimated, and what I see from the manuscript is that there is an effort to maintain claims that are not properly supported by the data and the final analysis. I recognized that I myself can also be in the wrong, so it will be left for eLife readers to make up their own minds.

We respectfully disagree with the assertion that our preregistered hypotheses "cannot be severely tested" due to the absence of predefined effect sizes and sample size estimation. First, the primary purpose of preregistration is to increase transparency and reduce analytical flexibility—not necessarily to formalize power calculations, especially when reliable effect size estimates are unavailable, as is the case in emerging areas of research such as this one (Nosek et al., 2018). Predefining effect sizes based on speculative or contextually irrelevant prior studies would have introduced unjustified assumptions and potentially distorted interpretation. In this case, we made the scientifically responsible choice to avoid anchoring our design on uncertain estimates.

Second, our sample size is comparable to or larger than those used in prior published studies addressing similar questions, and our analytic pipeline followed the preregistered plan. The reviewer suggests that we are attempting to maintain unsupported claims; however, our discussion is careful to reflect the limitations of the findings, including their statistical uncertainty, and we do not overstate conclusions. Where effects were weak or absent, we reported this transparently and interpreted them cautiously.

Finally, while formal power calculations are useful under certain conditions, they are not a prerequisite for a study to be considered a "severe test" of a hypothesis—especially when the statistical approach is preregistered, data are handled transparently, and null results are acknowledged rather than hidden. The sample size was informed by resource and feasibility constraints (Lakens, 2022) and aligns with or exceeds that of comparable studies in the field (Bowers et al., 2015). We have now made this rationale explicit in the manuscript.

“The sample size was informed by resource and feasibility constraints [98] and aligns with or exceeds that of comparable studies in the field [38], covering 60 participants in total, with two extra in case of incomplete data identified during later data analysis.”

We therefore maintain that the study is a valid and informative contribution, regardless of whether the observed results confirm or disconfirm the hypotheses.

Nonetheless, below we have revisited previous comments (#4, #5, #6, #7, #8 and #9) that reviewer 3 stated were not satisfactorily addressed.

Previous comment 4) Availability of the pre-registration

We consider this point now addressed, as access to the OSF pre-registration was resolved. Moreover, we can assure the reviewer that their identity was not revealed, as the request for access was not received/seen.

Previous comment 5) Magnitude of pain change – possibility of a type 1 error.

Previous comment 5) To be perfectly clear, I trust the results of this study more than some of the cited studies regarding nicotine and pain because it was preregistered, the sample size is considerably larger, and it seems carefully controlled. I just do not agree with the interpretation of the results, stated in the first paragraph of the Discussion. Quoting J. Cohen, "The primary product of a research inquiry is one or more measures of effect size, not P values" (Cohen, 1990). As I am sure the authors are aware of, even tiny differences between conditions, treatments or groups will eventually be statistically significant given arbitrarily large sample sizes. What really matters then is the magnitude of these differences. In general, the authors hypothesize on why there were no differences on the pressure pain model, and why decreases in heat pain were not mediated by PAF, but do not seem to consider the possibility that the intervention just did not cause the intended effect on the nociceptive system, which would be a much more straightforward explanations for all observations.

Our previous response included the fact that N=62 is not an arbitrarily large sample size and our consideration of the possibility of a false positive. In addition, we agree that conclusions should not be made based on p-values alone. Therefore, we have added more careful descriptions to highlight the nuances of our findings to the reader.

We edited the abstract and the first paragraph of the discussion to make clear to the reader that the effect sizes were small. In our desire to keep the text clear and readable, we did not intend to overstretch our claims, and apologise if the manuscript was unintentionally portraying this overstretch. These are very important sections of the manuscript in which to mention such nuances, as these are the parts viewed most when readers do not wish to read the entire manuscript.

Abstract

“The nicotine group showed a small decrease in heat-pain ratings compared to placebo group when controlling confounders, and a small increase in PAF across the scalp from pre- to post-gum, both with and without confounder adjustment. These effects were most pronounced in the central-parietal and right-frontal electrodes.”

Discussion first paragraph

“Compared to placebo, our data demonstrate three key findings: 1) nicotine increased PAF speed; 2) nicotine reduced prolonged heat pain intensity but only when controlling for confounding and not for the prolonged pressure pain model, and 3) decreases in prolonged heat pain intensity were not mediated by changes in PAF. This suggests that changes in heat pain due to nicotine occur through mechanisms unrelated to change in PAF and require replication as the effects required control for confounding and the reduction was small in magnitude.”

Discussion section on effect of nicotine on PAF:

“In line with previous literature in non-smokers [38,47,48], we contribute additional evidence that chewing 4 mg of nicotine gum produces a statistically significant increase in PAF speed in non-smoking, nicotine-gum-naïve, pain-free, adult participants. An exploratory cluster-based permutation analysis indicated that the change in wide-window PAF (8–12 Hz) produced by nicotine, was most pronounced in a cluster of electrodes in the central-parietal region and a separate cluster of electrodes in the right-frontal region, rather than in the pre-registered sensorimotor ROI (i.e., Cz, C1, C2, C3, C4).”

Discussion section on the effect of nicotine on pain:

“While acknowledging the modest effect size, it’s essential to consider the context of our study’s focus. Assessing the clinical relevance of pain reduction is pertinent in applications involving the use of any intervention for pain management [71]. However, from a mechanistic standpoint, particularly in understanding the implications of and relation to PAF, the specific magnitude of the pain effect becomes less pivotal, as small effects can still have mechanistic or biological meaning without establishing clinical meaning [72]. Nevertheless, future research should examine whether effects on pain increase in magnitude with different nicotine administration regimens (i.e. dose and frequency).”

Discussion section on relationship between PAF and pain:

“However, as this is the first study attempting to directly manipulate PAF, observing only small effects on pain, we cannot confidently state that a causal relationship between PAF and pain sensitivity does not exist, and future research should explore these relationships with larger sample sizes, larger doses of nicotine, and alternative PAF modulators.”

Discussion final paragraph:

“We have shown that, compared to placebo, nicotine gum produced statistically significant increases global PAF speed and decreases pain ratings during prolonged heat pain. However, the effects of nicotine on heat pain were not mediated by changes in PAF. Therefore, we find no evidence that PAF is causally related to pain. The implication of this is that changing PAF may not be a useful target for pain modulation, because pain sensitivity may not be dependent on state changes in PAF. However, this is the first study to test modulation of PAF as a mechanism for reducing pain intensity. Moreover, while effects on PAF appear small but robust, effects on heat pain required control for confounding to be observed. More randomised experiments on the effects of nicotine that use larger sample sizes, higher nicotine doses, and pharmacokinetic tracking, whilst also reducing mediator measurement error are needed to confirm the present findings and uncover whether these factors account for the lack of mediation effect seen here.”

Previous comment 6) Magnitude of pain change

Previous comment (6) In this regard, I do not believe that an average *increase* of 0.05 / 10 (Nicotine post pre) can be considered a "reduction of pain ratings", regardless of the contrast with placebo (average increase of 0.24 / 10). This tiny effect size is more relevant in the context of the considerable inter-individual variation, in which subjects scored the same heat pain model anywhere from 1 to 10, and the same pressure pain model anywhere from 1 to 8.5. In this regard, the minimum clinically or experimentally important differences (MID) in pain ratings varies from study to study and across painful conditions but is rarely below 1 / 10 in a VAS or NRS scale, see f. ex. (Olsen et al., 2017). It is not my intention to question whether nicotine can function as an acute analgesic in general (as stated in the Discussion), but instead, if it worked as such under these very specific experimental conditions. I also acknowledge that the authors note this issue in two lines in the Discussion, but I believe that this is not weighed properly.

Our previous response included discussion of interpreting mean values, a correction to a table that showed the ”increase” erroneously displayed previously due to list-wise deletion in the analysis. We have further attempted to weigh this issue more thoroughly and to always highlight the nuances of our results, as seen in our response to your previous comment 5 above. In further consideration of the effect size, note that the bootstrapped 95% CI for the effect of gum on change in pain was [-1.28, -0.16], which indicates that the overall difference falls within this range when controlling for other factors in the model. Of course this means the difference may be as small as -0.16, but also may be as large as -1.28. Such an effect might have be considered clinically relevant. Replication of the results is therefore warranted.

Also, as outlined in our response to reviewer 2 comment 1, our intention was not to establish a clinical utility at this stage. We hoped to achieve some biological/mechanistic meaning in the present study, and if our present findings are replicated in future, even these small-magnitude effects can shape our understanding of relevant mechanisms underlying pain and aid the direction of future clinical perspectives (Edenberg, 2021; Ross & Bassett, 2024). Meaningfulness in clinical (e.g., MID) or public policy contexts should of course be in mind, but were not the focus of this present work.

Several additional changes were made to the manuscript to emphasise these points to the reader, which are also outlined in the second round of response to your first comment 5 above, as both this comment 6 and 5 focus on magnitude of pain change. We hope this point is now considered to be weighed properly.

Previous comment 7) Magnitude of PAF change

Previous comment (7) In line with the topic of effect sizes, average effect sizes for PAF in the study cited in the manuscript range from around 1 Hz (Boord et al., 2008; Wydenkeller et al., 2009; Lim et al., 2016), to 2 Hz (Foulds et al., 1994), compared with changes of 0.06 Hz (Nicotine post - pre) or -0.01 Hz (Placebo post - pre). MIDs are not so clearly established for peak frequencies in EEG bands, but they should be certainly larger than some fractions of a Hertz (which is considerably below the reliability of the measurement).

In our first response to this comment, we acknowledged this difference in magnitude, agreed that MIDs for PAF are not established, discussed the potential impact of bin size, and highlighted this issue in the revised manuscript text. Although, note that we have now slightly reduced the emphasis on the bin size:

“The ability to detect changes in PAF could be considerably impacted by the frequency resolution used during Fourier Transformations, an element that is overlooked in recent methodological studies on PAF calculation [16,95].”

In addition, we would now like to highlight some more details regarding the prior studies.

We aimed to determine whether nicotine gum could increase PAF in healthy, pain-free, non-smokers. In line with previous literature in non-smokers (Foulds et al., 1994; Harkrider and Champlin, 2001; Bowers et al., 2015), we contribute additional evidence that chewing 4 mg of nicotine gum produces small increases in PAF speed in non-smoking, nicotine-gum-naïve, pain-free, adult participants. The other three papers mentioned by the reviewer are not cited in our manuscript and compare healthy individuals to chronic pain patients (Boord et al., 2008; Wydenkeller et al., 2009; Lim et al., 2016), and are thus not relevant to this discussion on the effect of nicotine on PAF.

A cluster-based permutation analysis indicated that the change in wide window PAF (8-12 Hz) produced by nicotine, was most pronounced in a cluster of electrodes in the central-parietal region and a separate cluster of electrodes in the right-frontal region, rather than in the pre-registered sensorimotor ROI (i.e., Cz, C1, C2, C3, C4). The mediation analysis results showed minimal differences when these clusters of electrodes were used compared to global results (global: b = 0.085, p < .018), except for the expected stronger effect of nicotine on change in PAF alongside greater confidence in this effect (bs = 0.11-0.14, ps < .003). See Supplementary Material 3.8-3.10.

Previous work examining nicotine-induced changes in PAF in healthy non-smokers have used small samples (i.e. <= 20) with nicotine injections or patches (Foulds et al., 1994; Harkrider and Champlin, 2001), while a later investigation used 62 male participants and nicotine gum (Bowers et al., 2015). Therefore, the present study is the largest study to date on nicotine and PAF including non-smoking male and female participants. Bowers didn’t see effects as large as 1 or 2 Hz (approx. 10 or 20% increases) seen in studies with fewer participants, such as n=4 (Foulds et al., 1994) and n=20 (Harkrider & Champlin, 2001), and the magnitude and location of effect depended on the COMT (Bowers et al., 2015).

Bowers et al (2015) showed that the Met allele was related to nicotine-induced increases in PAF, whereas Val/Val homozygotes did not demonstrate PAF changes. In addition, the distribution of PAF increases varied depending on whether Met allele carriers were heterozygotes (i.e. Val/Met: frontal, central, and occipital electrodes) or homozygotes (i.e. Met/Met: parietal electrode). The largest difference between placebo and nicotine gum was 0.51 Hz (a 5.8% increase) at the Fz electrode for the Val/Met group and the smallest significant difference was 0.34 Hz (a 3.5% increase) at the Pz electrode for the Met/Met group.

This study from Bowers used 6mg of nicotine gum, whereas we used 4mg, which may contribute to the lower effect size seen in our study. Dose-dependent effects have not been assessed in non-smokers, but given the dose-dependent changes seen in smokers who abstained from cigarettes prior to EEG recording (Pickworth et al., 1989; Pickworth et al., 1988; Pickworth et al., 1986), the smaller effect sizes seen in the present study may be in line with Bowers study.

Though genetic analysis was not included in the present work, the use of a lower nicotine dose, inclusion of male and female participants, use of more electrodes (64 rather than 8), and unmeasured genetic variations in PAF’s response to nicotine could explain the smaller effect sizes compared to (Bowers et al., 2015).

We have written a condensed summary of this to include in the manuscript discussion:

”Research on nicotine-induced changes in PAF in healthy non-smokers has largely relied on small samples (≤20) using injections or patches [47,48], with Bowers and colleagues [38] providing the only larger study (n= 62 males) using 6 mg nicotine gum. Their results showed modest PAF increases (~0.34–0.51 Hz; 3.5– 5.8%) compared to the larger effects (1–2 Hz; ~10–20%) reported in earlier small-sample studies, with magnitude and scalp distribution influenced by COMT genotype: Met allele carriers exhibited nicotineinduced increases (Val/Met: frontal, central, occipital; Met/Met: parietal), while Val/Val homozygotes showed no change. The present study now the largest to include both male and female non-smokers found smaller effects, likely due to the lower nicotine dose (4 mg), consistent with dose-dependent PAF modulation reported in abstinent smokers [69]. Broader EEG coverage (64 vs. 8 electrodes), sex inclusion, and unmeasured genetic variation may also account for differences between the present study and Bowers and colleagues [38]. Interestingly, beyond nicotine’s main effect on PAF, our 2W-LCS mediation model indicated a significant influence of sex, with females showing greater PAF increases than males. Given that female participants were significantly lighter than males (Supplementary Material 3.4) and the model did not test the interaction between sex at birth and gum group, these finding should be interpreted cautiously and warrant further investigation.”

Previous comment 8) removing confounders calls to question robustness of the findings. Simmons paper

Previous comment (8) The authors also ran alternative statistical models to analyze the data and did not find consistent results in terms of PHP ratings (PAF modulation was still statistically significantly different). The authors attribute this to the necessity of controlling for covariates. Now, considering the effects sizes, aren't these statistically significant differences just artifacts stemming from the inclusion of too many covariates (Simmons et al., 2011)? How much influence should be attributable to depression and anxiety symptoms, stress, sleep quality and past pain, considering that these are healthy volunteers? Should these contrasting differences call the authors to question the robustness of the findings (i.e., whether the same data subjected to different analysis provides the same results), particularly when the results do not align with the preregistered hypothesis (PAF modulation should occur on sensorimotor ROIs)?

In addition to our previous response, we would like to add the following: Although we were previously unaware of the author recommendations from Simmons et al (2011) during initial preparation of this manuscript, we find we cover all points and have added a few improvements to ensure our adherence to this guideline is clear.

We cover all six recommendations for authors suggested in Simmons et al (2011):

(1) We decided to collect 62 participants prior to commencing data collection, this covered 60 participants in total, with two extra in case of incomplete data identified during data analysis.

(2) We had more than 20 participants per group. We have now outlined this decision in the methods section:

“The sample size was informed by resource and feasibility constraints [98] and aligns with or exceeds that of comparable studies in the field [38], covering 60 participants in total, with two extra in case of incomplete data identified during later data analysis.”

(3) We provide detailed methods, which outlines all the variables collected.

(4) We report all experimental conditions, and had no failed manipulations to report. We have now slightly adjusted the wording of our experimental protocol within the methods section to make this clear:

“Participants were randomly assigned to one of only two conditions, thus after the second resting-state EEG, either nicotine or placebo gum was administered with the chew procedure. Following completion of the chew procedure, the third resting state, the second pain assessments, and the fouth resting state were conducted.”

Regarding recommendation (5), listwise deletion was used, as already stated in the manuscript (see Methods, Statistical analysis, Mediation analysis, Assumptions), to exclude participants for whom only one timepoint was available in either the pain or PAF measures, and thus it is not possible to conduct analysis on participants with this missing data. Otherwise, no participants were removed for other reasons. 6) As we conducted analyses using pre-registered covariates determined based on a DAG, we repeated analyses without including these covariates in the form of RM-ANOVAs. The effect on PAF remains, however, as Simmons et al (2011) suggests, the fact we no longer observed a significant effect of gum on heat pain ratings may indicate that the finding is reliant on the presence of a covariate, which does not necessarily mean that the finding is not relevant. However, as reviewer 3 suggests, this should also cause us to question the robustness of the findings and the possibility of type 1 error. Therefore, we have adjusted text in the discussion as follows:

“This underscores the necessity of controlling for confounding variables, as nicotine may only produce pain reduction for those with higher stress, and calls for further investigation into the robustness of the effect of nicotine on heat pain ratings.”

Lastly, regarding the influence that should be attributable to our confounding health-related measures, given the healthy population, we still observe variation in several of these confounders. ’Healthy’ populations still score across a range on these scales, with cut-offs acting as suggestions for problematic symptomology. Assumptions testing confirmed the normality of these covariates, and lack of collinearity so that they could be included in our analyses. We have added a sentence to the Sample Characteristics in Results to outline the ranges reported for the confounding variables:

“Depressive symptoms ranged from 0-4 (out of a maximum of 6), anxiety from 0-6 (out of a maximum of 6), perceived stress from 2-18 (out of a maximum of 40), and sleep quality from 2-10 (out of a maximum of 10; i.e. “excellent sleep quality”).”

Previous comment 9) Exploratory sex analysis?

(9) Beyond that, I believe in some cases that the authors overreach in an attempt to provide explanations for their results. While I agree that sex might be a relevant covariate, I cannot say whether the authors are confirming a pre-registered hypothesis regarding the gender-specific correlation of PAF and pain, or if this is just a post hoc subgroup analysis. Given the large number of analyses performed (considering the main document and the supplementary files), caution should be exercised on the selective interpretation of those that align with the researchers' hypotheses.

To reiterate our previous response to this comment, we stated that sex was an exploratory analysis in the results and was chosen based on other PAF-pain literature, and had also adjusted part of the discussion to align with this. However, we have now seen the first paragraph of the discussion was not adjusted. Therefore, we hope the below edit has now adequately address the previous reviewer 3 comment 9:

“In addition, when assessing the whole sample, we found inconclusive evidence for 4) slower PAF speed at baseline being associated with higher pain ratings for both the heat and pressure pain models. However, during exploratory analysis we found moderate evidence that slower PAF speed at baseline is associated with higher heat pain ratings at baseline in males but not females. In contrast PAF was not associated with pressure pain ratings for males or females.”

References for response to reviewer 3

Boord P, Siddall PJ, Tran Y, Herbert D, Middleton J, Craig A. Electroencephalographic slowing and reduced reactivity in neuropathic pain following spinal cord injury. Spinal cord. 2008 Feb;46(2):118-23.

Bowers H, Smith D, De La Salle S, Choueiry J, Impey D, Philippe T, et al. COMT polymorphism modulates the resting-state EEG alpha oscillatory response to acute nicotine in male non-smokers. Genes, Brain and Behavior 2015;14:466–76.

Edenberg HJ. Perspective on beyond statistical significance: Finding meaningful effects. Complex Psychiatry 2021;7:1–8. https://doi.org/10.1159/000517237.

Foulds J, McSorley K, Sneddon J, Feyerabend C, Jarvis MJ, Russell MA. Effect of subcutaneous nicotine injections on EEG alpha frequency in non-smokers: A placebo-controlled pilot study. Psychopharmacology 1994;115:163–6.

Harkrider AW, Champlin CA. Acute effect of nicotine on non-smokers: III. LLRs and EEGs. Hearing Research 2001;160:99–110.

Lakens D. Sample size justification. Collabra: psychology. 2022 Mar 22;8(1):33267.

Lim M, Kim JS, Kim DJ, Chung CK. Increased low-and high-frequency oscillatory activity in the prefrontal cortex of fibromyalgia patients. Frontiers in human neuroscience. 2016 Mar 14;10:111.

Nosek BA, Ebersole CR, DeHaven AC, Mellor DT. The preregistration revolution. Proceedings of the National Academy of Sciences. 2018 Mar 13;115(11):2600-6.

Pickworth WB, Herning RI, Henningfield JE. Electroencephalographic effects of nicotine chewing gum in humans. Pharmacology Biochemistry and Behavior 1986;25:879–82.

Pickworth WB, Herning RI, Henningfield JE. Mecamylamine reduces some EEG effects of nicotine chewing gum in humans. Pharmacology, Biochemistry, and Behavior 1988;30:149–53.

Pickworth WB, Herning RI, Henningfield JE. Spontaneous EEG changes during tobacco abstinence and nicotine substitution in human volunteers. Journal of Pharmacology and Experimental Therapeutics 1989;251:976–82.

Ross LN, Bassett DS. Causation in neuroscience: keeping mechanism meaningful. Nature Reviews Neuroscience. 2024 Feb;25(2):81-90.

Simmons JP, Nelson LD, Simonsohn U. False-positive psychology: Undisclosed flexibility in data collection and analysis allows presenting anything as significant. Psychological science. 2011 Nov;22(11):1359-66.

Wydenkeller S, Maurizio S, Dietz V, Halder P. Neuropathic pain in spinal cord injury: significance of clinical and electrophysiological measures. European Journal of Neuroscience. 2009 Jul;30(1):91-9.

Changes from the earlier version to the revised version

Abstract

Earlier version: Recent research has linked individual peak alpha frequency (PAF) to pain sensitivity, but whether PAF alterations can influence pain remains unclear. Our study investigated the effects of nicotine on pain sensitivity and whether pain changes are mediated by PAF changes. In a randomised, double-blind, placebo-controlled experiment, 62 healthy adults (18–44 years) received either 4 mg nicotine gum (n=29) or placebo (n=33). Resting state EEG and pain ratings during prolonged heat and pressure models were collected before and after nicotine intake. The nicotine group had a decrease in heat pain ratings compared to the placebo group and increased PAF speed across the scalp from pre to post-gum, driven by changes at central-parietal and right-frontal regions. However, mediation analysis did not support the notion that PAF changes mediate nicotine’s effects on pain sensitivity. While a growing body of literature supports a link between PAF and both acute and chronic pain, further work is needed to understand the mechanisms of this link.

Revised version: Recent research has linked individual peak alpha frequency (PAF) to pain sensitivity, but whether PAF alterations can influence pain remains unclear. Our study investigated the effects of nicotine on pain sensitivity and whether pain changes are mediated by PAF changes. In a randomised, double-blind, placebo-controlled experiment, 62 healthy adults (18–44 years) received either 4 mg nicotine (n=29) or placebo gum (n=33). Resting state EEG and pain ratings during prolonged heat and pressure models were collected before and after nicotine intake. The nicotine group showed a small decrease in heat-pain ratings compared to placebo group when controlling confounders, and a small increase in PAF across the scalp from pre- to post-gum, both with and without confounder adjustment. These effects were most pronounced in the central-parietal and right-frontal electrodes. However, mediation analysis did not support the notion that PAF changes mediate nicotine’s effects on pain sensitivity. While a growing body of literature supports a link between PAF and both acute and chronic pain, further work is needed to understand the mechanisms of this link.

Introduction

Chronic pain is a highly prevalent and multifaceted health issue, which produces significant physical, emotional, and economic burden to those in pain, their families and friends, health care systems, and society [1–3]. Identifying modifiable factors that influence pain sensitivity could be a key step in reducing the presence and burden of chronic pain [4–6]. Individual peak alpha frequency (PAF) is an electro-physiological brain measure that shows promise as a biomarker of pain sensitivity [7–13]. Although the number of studies investigating PAF as a predictor of pain is increasing [14–18], no experiments have manipulated PAF while also measuring pain, which has limited our understanding of whether, and how, PAF is mechanistically linked to pain sensitivity.

PAF refers to the dominant oscillatory frequency in the 8–12 Hz alpha range. As alpha oscillations are believed to represent periodic inhibitory processes that direct temporal aspects of sensory information processing [19,20], it is thought that PAF reflects the rate at which sensory information is sampled from the external environment [21–23]. In the context of pain, slower PAF speed has been associated with the presence and duration of chronic pain [17,24–28], while in healthy, pain-free individuals, slower PAF has been associated with higher pain sensitivity in experimental pain models [8–10] and acute post-surgical settings [11].

Earlier version: While previous studies suggest that PAF could be used as either a biomarker for disease progression or as a predictive biomarker for pain sensitivity, critical gaps remain in our understanding of how PAF is mechanistically linked to pain. Specifically, no study has examined whether the manipulation of PAF alters pain sensitivity. Addressing this question has significant implications, as it would determine i) whether PAF is causally implicated in pain processing, and ii) whether manipulation of PAF has relevance as a method of clinical pain modulation. Several methods have been used to alter individual PAF, including brain stimulation [29–31], exercise [32–34], visual stimulation [35,36], cognitive tasks [37], and nicotine [38–40]. Because evidence suggests that nicotine can modulate PAF, where both nicotine and smoking increase PAF speed [38,41–48], we chose nicotine to assess our aim of whether change in PAF mediate changes in pain in a ‘mediation by design’ approach [49]. In addition, given evidence that nicotine

Revised version: While previous studies suggest that PAF could be used as either a biomarker for disease progression or as a predictive biomarker for pain sensitivity, critical gaps remain in our understanding of how PAF is mechanistically linked to pain. Specifically, no study has examined whether the manipulation of PAF alters pain sensitivity. Addressing this question has significant implications, as it would determine i) whether PAF is causally implicated in pain processing, and ii) whether manipulation of PAF has relevance as a method of clinical pain modulation. Several methods have been used to alter individual PAF, including brain stimulation [29–31], exercise [32–34], visual stimulation [35,36], cognitive tasks [37], and nicotine [38–40]. Because evidence suggests that nicotine can modulate PAF, where both nicotine and smoking increase PAF speed [38,41–48], we chose nicotine to assess our aim of whether changes in PAF mediate changes in pain in a ‘mediation by design’ approach [49]. In addition, given evidence that nicotine may increase experimental pain thresholds and tolerance [50–54], nicotine could also influence pain ratings during tonic pain. Thus, the primary questions addressed in this study were whether nicotine can: i) increase peak alpha frequency (PAF), ii) reduce pain sensitivity, and iii) determine whether changes in pain sensitivity are mediated by changes in PAF speed.

Removed in the revised version: may increase experimental pain thresholds and tolerance [50–54], nicotine could also influence pain ratings during tonic pain. Thus, the primary questions addressed in this study were whether nicotine can: i) increase peak alpha frequency (PAF), ii) reduce pain sensitivity, and iii) determine whether changes in pain sensitivity are mediated by changes in PAF speed.

A secondary question pertains to the nature of the PAF–pain relationship. A substantial proportion of the current literature suggests that slower PAF speed is associated with higher pain sensitivity [8–11,17,24–27]. However, conflicting findings exist, with some studies reporting no relationship between PAF and chronic pain [55–58], or an opposite relationship where faster PAF is associated with higher pain sensitivity in experimental pain settings [7,15]. Divergent results could be due to variations in the relationship between PAF and different types of pain, or the dependence of this relationship on individual factors. For example, it was recently suggested that PAF–pain relationships may only be observed in neuropathic pain or expressed differently in males and females [17]. One way to address this issue is to investigate PAF across multiple pain models [12]; currently the PAF–pain relationship has been studied using thermal [7,8,10], musculoskeletal [9,15], and surgical [11] pain, but never with multiple pain modalities included in the same study. Thus, a secondary question was whether baseline PAF, measured in a pain-free resting state, correlated with pain sensitivity during a model of heat pain (i.e. PHP), which has been previously assessed and is negatively correlated with PAF [10], and a model of pressure pain (i.e. cuff pressure algometry; CPA [59–63]) that has not been assessed in relation to PAF. Applying two experimental pain models also allows assessment of the generalisability of nicotine’s effects on pain and the potential mediating effects of PAF, including whether these effects are consistent across different types of pain.

Removed in the revised version: Using a randomised, placebo-controlled, double-blind, parallel design, a sample of healthy, pain-free, non-smoking participants had their resting state PAF measured before and after two models of prolonged pain (i.e. CPA and PHP; pre-gum). Participants then chewed either nicotine or placebo gum, after which their resting state PAF was re-assessed before and after the two prolonged pain models (i.e. CPA and PHP; post-gum). We predicted that: 1. nicotine gum would increase PAF speed in

Earlier version: comparison to placebo gum; 2. nicotine gum would decrease pain ratings in comparison to placebo gum; 3. decreases in pain ratings following nicotine gum would be mediated by increases in the speed of PAF; and 4. faster baseline PAF would be associated with lower baseline pain ratings. Based on a pre-registered analysis plan (https://osf.io/pc4rq/), predictions 1–3 were tested using two-wave latent change score (2W-LCS) mediation models [64–66], while prediction 4 was tested using correlation analysis.

Revised version: Using a randomised, placebo-controlled, double-blind, parallel design, a sample of healthy, pain-free, non-smoking participants had their resting state PAF measured before and after two models of prolonged pain (i.e. CPA and PHP; pre-gum). Participants then chewed either nicotine or placebo gum, after which their resting state PAF was re-assessed before and after the two prolonged pain models (i.e. CPA and PHP; post-gum). We predicted that: 1. nicotine gum would increase PAF speed in comparison to placebo gum; 2. nicotine gum would decrease pain ratings in comparison to placebo gum; 3. decreases in pain ratings following nicotine gum would be mediated by increases in the speed of PAF; and 4. faster baseline PAF would be associated with lower baseline pain ratings. Based on a pre-registered analysis plan (https://osf.io/pc4rq/), predictions 1–3 were tested using two-wave latent change score (2W-LCS) mediation models [64–66], while prediction 4 was tested using correlation analysis.

Results

The statistical analysis plan for this study was pre-registered on Open Science Framework (OSF: https://osf.io/pc4rq/) in July 2022.

Sample Characteristics

The participant sample included 62 healthy adults (26.45 ± 7.28; range: 18–44 years), without neurological or psychological disorders. Fourteen participants self-identified as North-West European (22.58%), 12 as South-East Asian, (19.36%), and eight reported mixed ancestries (12.90%; see Supplementary Material 1 for a full list of self-reported ancestry). There were 32 females (24.53 ± 6.52; range: 18–43 years) and 30 males (28.5 ± 7.58; range: 19–44 years); no participants reported their sex at birth as inter-sex. Thirty-one individuals identified as women, 30 as men, and 1 as non-binary.

Earlier version: Regarding pain in the past two weeks, 20 participants reported “None” (nicotine: n = 10; placebo: n = 10), 16 reported “Very Mild” (nicotine: n = 14; placebo: n = 12), 15 reported “Mild” (nicotine: n = 5; placebo: n = 10), one reported “Moderate” (nicotine: n = 1, placebo: n = 0), and no participants reported “Severe” pain in the past two weeks. Participants were successfully blinded to group allocation (Supplementary Material 2) and the remaining baseline demographics are displayed in Table 1. Following CONSORT recommendations [67,68], no baseline statistics are reported.

Revised version: Regarding pain in the past two weeks, 20 participants reported “None” (nicotine: n = 10; placebo: n = 10), 16 reported “Very Mild” (nicotine: n = 14; placebo: n = 12), 15 reported “Mild” (nicotine: n = 5; placebo: n = 10), one reported “Moderate” (nicotine: n = 1, placebo: n = 0), and no participants reported “Severe” pain in the past two weeks. Participants were successfully blinded to group allocation (Supplementary Material 2) and baseline demographics are displayed in Table 1. Depressive symptoms ranged from 0–4 (out of a maximum of 6), anxiety from 0–6 (out of a maximum of 6), perceived stress from 2–18 (out of a maximum of 40), and sleep quality from 2–10 (out of a maximum of 10; i.e. “excellent sleep quality”). Following CONSORT recommendations [67,68], no baseline statistics are reported.

Effects of nicotine gum on alpha oscillations and prolonged pain ratings

Mediation analysis was conducted to assess the effects of nicotine on PAF and pain ratings, and whether change in PAF mediated change in pain.

ANCOVA-equivalent two-wave latent change score (2W-LCS) models (i.e. with autoregressive and coupling parameters) and difference score 2W-LCS models (i.e. without auto-regressive and coupling parameters) were constructed for mean PHP and CPA pain outcomes separately (see Methods for further details). As no significant difference was found between model fits (see Supplementary Materials 3.1 and 3.2), the difference score 2W-LCS models are reported here; whilst the full ANCOVA-equivalent 2W-LCS models are presented in Supplementary Material 3.1.

Nicotine gum modulated global PAF (8–12 Hz) and PHP ratings, but change in PHP ratings were not mediated by change in PAF

The PHP model consisted of five repetitions of a 40 second heat stimulation to the inside left forearm at a fixed 46 °C (see Methods), which was most frequently described as moderate/severe hot, sharp, and shooting pain (Figure 1, plot A). Raw time courses of pain ratings during PHP are displayed in Figure 1 (plot C), from which the mean pain during max heat stimulation was calculated (Table 2, Figure 1, plot E).

Earlier version: Percentage of participants reporting each category descriptor of pain on the McGill pain questionnaire for A) phasic heat pain (PHP) and B) cuff pressure algometry (CPA) at baseline. Average pain time courses split into nicotine (orange, n = 29) and placebo (grey, n = 33) gum groups for C) PHP and D) CPA models of prolonged pain. Raincloud plots of mean pain ratings during E) PHP and F) CPA models of prolonged pain.

Revised version: Percentage of participants reporting each category descriptor of pain on the McGill pain questionnaire for A) phasic heat pain (PHP) and B) cuff pressure algometry (CPA) at baseline.

Added in the revised version: Average pain time courses split into nicotine (orange, n = 29) and placebo (grey, n = 33) gum groups for C) PHP and D) CPA models of prolonged pain. Raincloud plots of mean pain ratings during E) PHP and F) CPA models of prolonged pain.

Five participants could not tolerate the fixed 46 °C stimulus during familiarisation, therefore a lower temperature of 45 °C was used for both pre- and post-gum PHP assessments for these participants, as reported in the protocol (see Methods). Two participants were excluded from mediation analysis assessing the effects of nicotine on PHP ratings: one from the placebo group due to an equipment issue during the second pain assessment, and one from the nicotine group that could not tolerate the baseline temperature for the second pain assessment (nicotine: n = 28; placebo: n = 32).

Earlier version: The difference score model for PHP (Figure 3; model fit: χ2(24) = 25.874, p = 0.36, CFI = 0.99, RMSEA = 0.036, SRMR = 0.078) suggests a significant total (b =−0.62, p = .023, bootstrapped 95% CI: [-1.21, -0.15]) and direct (b = −0.68, p = .021, bootstrapped 95% CI: [-1.28, -0.16]) effect of nicotine on change in PHP ratings (Figure 1). Change in pain was reduced by a factor of -0.667 for the nicotine group compared to the placebo group when controlling for confounding variables (Figure 3).

Revised version: The difference score model for PHP (Figure 3; model fit: 𝜒2(24) = 25.874, 𝑝 = 0.36, CFI = 0.99, RMSEA = 0.036, SRMR = 0.078) suggests a significant total (𝑏 = −0.62, 𝑝 = .023, bootstrapped 95% CI: [−1.21, −0.15]) and direct (𝑏 = −0.68, 𝑝 = .021, bootstrapped 95% CI: [−1.28, −0.16]) effect of nicotine on change in PHP ratings (Figure 1). Change in pain was reduced by a factor of −0.68 for the nicotine group compared to the placebo group when controlling for confounding variables (Figure 3). See Supplementary Material 3.2, Table 5, for full model output.

Earlier version: Frequency spectra (plots A and B) and topographical plots (plot E) of PAF pre-post nicotine, are displayed by group in Figure 2, with PAF values displayed in Table 3. The difference score model suggests a significant effect of nicotine gum on change in global PAF (b = 0.085, p = .018, bootstrapped 95% CI: [0.017, 0.15]), suggesting change in PAF was increased by a factor of 0.085 for the nicotine group compared to placebo when controlling for confounding variables (Figure 3). Therefore, nicotine changed PAF and PHP ratings in the predicted directions compared to what would be expected without nicotine (i.e. placebo). However, as the estimated effect of change in PAF speed on change in PHP ratings was non-significant (b = 0.23, p = .78, bootstrapped 95% CI: [-1.28, 1.9]), the change in PAF did not mediate the change in PHP ratings (i.e. indirect effect; b = 0.019, p = .81, bootstrapped 95% CI: [-0.093, 0.23]).

Revised version: Frequency spectra (plots A and B) and topographical plots (plot E) of PAF pre-post nicotine, are displayed by group in Figure 2, with PAF values displayed in Table 3. The difference score model suggests a significant effect of nicotine gum on change in global PAF (𝑏 = 0.085, p = .018, bootstrapped 95% CI: [0.017, 0.15]), suggesting change in PAF was increased by a factor of 0.085 for the nicotine group compared to placebo when controlling for confounding variables (Figure 3). Therefore, nicotine changed PAF and PHP ratings in the predicted directions compared to what would be expected without nicotine (i.e. placebo). However, as the estimated effect of change in PAF speed on change in PHP ratings was non-significant (𝑏 = 0.23, 𝑝 = .78, bootstrapped 95% CI: [-1.28, 1.9]), the change in PAF did not mediate the change in PHP ratings (i.e. indirect effect; 𝑏 = 0.019, 𝑝 = .81, bootstrapped 95% CI: [-0.093, 0.23]).

Earlier version: A) Individual spectral plots of eyes closed (EC) resting state EEG data showing dominant peaks within the alpha range of 8–12 Hz. Note that many participants’ peaks fall outside the narrow 9–11 Hz window. Top panel shows individual spectra for participants in the placebo group (in grey, n = 33), bottom panel shows individual spectra for participants in the nicotine group (in orange, n = 29). B) Spectral plots of eyes closed (EC) resting state EEG shows averaged spectra by group and time point. C) Raincloud plots of mean global and D) sensorimotor PAF values, calculated with a wide window (8–12 Hz) and the centre of gravity method, before and after chewing either nicotine (in orange) or placebo (in grey) gum. E) Topographical plots of PAF calculated for each electrode with the centre of gravity method and 8–12 Hz window are shown for the pre-gum baseline (rest 1) and post-gum (rest 3). F) Topographical plot of t-test statistics produced by cluster-based permutation analysis of nicotine group, displaying significant difference in PAF after chewing nicotine gum in two separate clusters marked by crosses (x).

Revised version: A) Individual spectral plots of eyes closed (EC) resting state EEG data showing dominant peaks within the alpha range of 8–12 Hz. Note that many participants’ peaks fall outside the narrow 9–11 Hz window. Top panel shows individual spectra for participants in the placebo group (in grey, n = 33), bottom panel shows individual spectra for participants in the nicotine group (in orange, n = 29). B) Spectral plots of eyes closed (EC) resting state EEG shows averaged spectra by group and time point. C) Raincloud plots of mean global and D) sensorimotor PAF values, calculated with a wide window (8–12 Hz) and the centre of gravity method, before and after chewing either nicotine (in orange) or placebo (in grey) gum. E) Topographical plots of PAF calculated for each electrode with the centre of gravity method and 8–12 Hz window are shown for the pre-gum baseline (rest 1) and post-gum (rest 3). F) Topographical plot of t-test statistics produced by cluster-based permutation analysis of nicotine group, displaying significant difference in PAF after chewing nicotine gum in two separate clusters marked by crosses where p<.01 for adjacent electrodes (x). These effects were most pronounced in the central-parietal and right-frontal electrodes.

Earlier version: Path model displaying unstandardised estimates and p-values for the difference score two-wave latent change score (2W-LCS) model assessing the effect of nicotine against placebo (i.e. gum group) on change in global 8–12 Hz peak alpha frequency (i.e. ΔPAF) and change in mean phasic heat pain (PHP) ratings (i.e. ΔPain). Green arrows highlight the indirect effects (i.e. a and b), the red arrow highlights the direct effect (i.e. c’), the blue lines highlight pre-determined confounding variables, and dashed lines denote fixed effects. This model suggests that chewing nicotine gum, compared to placebo, increases the speed of PAF, and decreases PHP ratings, however the change in PAF does not influence the change in PHP ratings, therefore changes in PHP ratings after nicotine gum are not mediated by changes in PAF.

Revised version: Path model displaying unstandardised estimates and p-values for the difference score two-wave latent change score (2W-LCS) model assessing the effect of nicotine against placebo (i.e. gum group) on change in global 8–12 Hz peak alpha frequency (i.e. Δ𝑃𝐴𝜞) and change in mean phasic heat pain (PHP) ratings (i.e. Δ𝑃𝑎𝑖𝑛). Green arrows highlight the indirect effects (i.e. a and b), the red arrow highlights the direct effect (i.e. c’), the blue lines highlight pre-determined confounding variables, and dashed lines denote fixed effects. This model suggests that chewing nicotine gum, compared to placebo, increases the speed of PAF, and decreases PHP ratings, however the change in PAF does not influence the change in PHP ratings, therefore changes in PHP ratings after nicotine gum are not mediated by changes in PAF.

Earlier version: There was also a significant effect of perceived stress at baseline on change in PHP ratings when controlling for group allocation and other confounding variables (b =−0.096, p = .048, bootstrapped 95% CI: [-0.19, -0.000047]), where higher perceived stress resulted in larger decreases in PHP ratings (see Supplementary Material 3.3 for post-hoc analysis of stress). There was also a significant effect of sex on change in PAF (b = -0.077, p = .037, bootstrapped 95% CI: [-0.16, -0.011]); change in PAF was reduced by a factor of -0.077 for males compared to females when controlling for gum group and confounders, therefore the increase in PAF speed was greater for females than males (see Supplementary Material 3.4 for post-hoc analysis of sex differences).

Revised version: There was also a significant effect of perceived stress at baseline on change in PHP ratings when controlling for group allocation and other confounding variables (𝑏 = −0.096, 𝑝 = .048, bootstrapped 95% CI: [−0.19, −0.000047]), where higher perceived stress resulted in larger decreases in PHP ratings (see Supplementary Material 3.3 for post-hoc analysis of stress). There was also a significant effect of sex on change in PAF (𝑏 = −0.077, p = .037, bootstrapped 95% CI: [−0.16, −0.011]); change in PAF was reduced by a factor of −0.077 for males compared to females when controlling for gum group and confounders, therefore the increase in PAF speed was greater for females than males (see Supplementary Material 3.4 for post-hoc analysis of sex differences).

Earlier version: A complete table of PAF and power values for all four time points is displayed in Supplementary Material 4.1. Exploratory analysis of interactions between time point (i.e. pre-/post-gum) and group (i.e. nicotine/placebo) on PHP ratings, PAF speed, PHP ratings, and alpha power, using repeated measures analysis of variance (RM-ANOVA) that do not control for confounding, are displayed in Supplementary Materials 4 and Notably, exploratory RM-ANOVAs still indicated interactions between gum group and time point on PAF (8–12 Hz) indicating that PAF only increased for the nicotine group after chewing (Supplementary Material 4), but there was no interaction between gum group and time point on PHP ratings (Supplementary Material 5), underscoring the necessity of controlling for confounding variables.

Revised version: A complete table of PAF and power values for all four time points is displayed in Supplementary Material 4.1. Exploratory analysis of interactions between time point (i.e. pre-/post-gum) and group (i.e. nicotine/placebo) on PHP ratings, PAF speed, PHP ratings, and alpha power, using repeated measures analysis of variance (RM-ANOVA) that do not control for confounding, are displayed in Supplementary Materials 4 and 5. Notably, exploratory RM-ANOVAs still indicated interactions between gum group and time point on PAF (8–12 Hz) indicating that PAF only increased for the nicotine group after chewing (Supplementary Material 4), but there was no interaction between gum group and time point on PHP ratings (Supplementary Material 5). This underscores the necessity of controlling for confounding variables, as nicotine may only produce pain reduction for those with higher stress, and calls for further investigation into the robustness of the effect of nicotine on heat pain ratings.

Nicotine gum modulated global PAF (8–12 Hz), with no total, direct, or indirect effects on CPA ratings

The CPA model involved four minutes of continuous cuff pressure to the left calf muscle that was initially calibrated to each individual’s 4/10 pain (see Methods), and was most frequently described as mild/moderate cramping, aching, throbbing, and heavy pain (Figure 1, plot B). Raw time courses of pain ratings during CPA are displayed in Figure 1 (plot D), from which mean pain during the 4 minutes of stimulation was calculated (Table 2, Figure 1, plot F). No participants were excluded from the analyses of pressure pain (nicotine: n = 29; placebo: n = 33).

Earlier version: Estimates from the difference score 2W-LCS model for CPA pain ratings (Figure 4) for total (b = −0.28, p = .50, bootstrapped 95% CI: [-1.13, 0.49]), direct (b = −0.17, p = .69, bootstrapped 95% CI: [-1.07, 0.59]), and indirect (b = −0.11, p = .53, bootstrapped 95% CI: [-0.61, 0.18]) effects of nicotine on CPA ratings were non-significant (model fit: χ2(24) = 29.55, p = 0.20, CFI = 0.97, RMSEA = 0.061, SRMR = 0.085). Therefore, there was no evidence to suggest that nicotine gum directly or indirectly influenced change in CPA ratings.

Revised version: Estimates from the difference score 2W-LCS model for CPA pain ratings (Figure 4) for total (𝑏 = −0.28, 𝑝 = .50, bootstrapped 95% CI: [-1.13, 0.49]), direct (𝑏 = −0.17, 𝑝 = .69, bootstrapped 95% CI: [-1.07, 0.59]), and indirect (𝑏 = −0.11, 𝑝 = .53, boot-strapped 95% CI: [-0.61, 0.18]) effects of nicotine on CPA ratings were non-significant (model fit: 𝜒2(24) = 29.55, 𝑝 = 0.20, CFI = 0.97, RMSEA = 0.061, SRMR = 0.085). Therefore, there was no evidence to suggest that nicotine gum directly or indirectly influenced change in CPA ratings.

Earlier version: Path model displaying unstandardised estimates and p-values for the difference score two-wave latent change score (2W-LCS) model assessing the effect of nicotine against placebo (i.e. gum group) on change in global 8–12 Hz peak alpha frequency (i.e. ΔPAF) and change in mean cuff pressure algometry (CPA) ratings (i.e. ΔPain). Green arrows highlight the indirect effects (i.e. a and b), the red arrow highlights the direct effect (i.e. c’), the blue lines highlight pre-determined confounding variables, and dashed lines denote fixed effects. This model suggests that chewing nicotine gum, compared to placebo, increases the speed of PAF and does not influence CPA ratings, directly or indirectly.

Revised version: Path model displaying unstandardised estimates and p-values for the difference score two-wave latent change score (2W-LCS) model assessing the effect of nicotine against placebo (i.e. gum group) on change in global 8–12 Hz peak alpha frequency (i.e. Δ𝑃𝐴𝜞) and change in mean cuff pressure algometry (CPA) ratings (i.e. Δ𝑃𝑎𝑖𝑛). Green arrows highlight the indirect effects (i.e. a and b), the red arrow highlights the direct effect (i.e. c’), the blue lines highlight pre-determined confounding variables, and dashed lines denote fixed effects. This model suggests that chewing nicotine gum, compared to placebo, increases the speed of PAF and does not influence CPA ratings, directly or indirectly.

Earlier version: Consistent with the PHP model, there was a significant effect of nicotine gum on change in global PAF (8–12 Hz) speed (b = 0.089, p = .011, bootstrapped 95% CI: [0.018, 0.16]), where change in PAF speed increased by 0.089 for the nicotine group compared to the placebo group. Additionally, there was again an effect of sex on change in PAF (b = -0.076, p = .027, bootstrapped 95% CI: [-0.15, -0.01]), where the increase in PAF speed was greater for females than males. And the estimated effect of change in PAF speed on change in pain ratings was non-significant for CPA (b = −1.28, p = .49, bootstrapped 95% CI: [-5.21, 2.11]).

Revised version: Consistent with the PHP model, there was a significant effect of nicotine gum on change in global PAF (8–12 Hz) speed (𝑏 = 0.089, p = .011, bootstrapped 95% CI: [0.018, 0.16]), where change in PAF speed increased by 0.089 for the nicotine group compared to the placebo group. Additionally, there was again an effect of sex on change in PAF (𝑏 = −0.076, p = .027, bootstrapped 95% CI: [−0.15, −0.01]), where the increase in PAF speed was greater for females than males. And the estimated effect of change in PAF speed on change in pain ratings was non-significant for CPA (𝑏 = −1.28, 𝑝 = .49, bootstrapped 95% CI: [-5.21, 2.11]).

Exploratory analysis of interactions between time point (i.e. pre-/post-gum) and group (i.e. nicotine/placebo) on CPA ratings using RM-ANOVA (Supplementary Material 5) suggested that CPA pain ratings decreased after chewing gum irrespective of group (Figure 1, plots D and F).

Added in the revised version: Exploratory data-driven analysis of alpha

Removed in the revised version: Exploratory cluster-based permutation analysis

The effect of nicotine on PAF speed was only significant for global PAF with a wide frequency window (8–12 Hz) and not for sensorimotor PAF (i.e. Cz, C1, C2, C3, C4) or narrow frequency windows (9–11 Hz; Supplementary Materials 3.4–3.6). This suggests that nicotine produces change in PAF somewhere in sensor-space, but not in the pre-defined sensorimotor region comprising central electrodes. An exploratory cluster-based permutation analysis (see Methods) showed significant increases in PAF speed after chewing nicotine gum, with this effect being maximal over two spatially separated clusters of electrodes: one cluster over the central-parietal region, and another over the right-frontal region (Figure 2, plot F). None of the electrodes exceeded the threshold of .05 to be clustered in the placebo group.

Earlier version: The potential mediating effect of this change in PAF on change in PHP and CPA was explored (not pre-registered) by averaging within each cluster (central-parietal: CP1, CP2, Cpz, P1, P2, P3, P4, Pz, POz; right-frontal: F8, FT8, FT10) and across both clusters. This averaging across electrodes produced three new variables, each assessed in relation to mediating effects on PHP and CPA ratings. The resulting in six exploratory mediation analysis (difference score 2W-LCS) models demonstrated minimal differences from the main analysis of global PAF (8-12 Hz), except for the expected stronger effect of nicotine on change in PAF (bs = 0.11–0.14, ps < .003; Supplementary Materials 3.8–3.10).

Revised version: The potential mediating effect of this change in PAF on change in PHP and CPA was explored (not pre-registered) by averaging within each cluster (central-parietal: CP1, CP2, Cpz, P1, P2, P3, P4, Pz, POz; right-frontal: F8, FT8, FT10) and across both clusters. This averaging across electrodes produced three new variables, each assessed in relation to mediating effects on PHP and CPA ratings. The resulting in six exploratory mediation analysis (difference score 2W-LCS) models demonstrated minimal differences from the main analysis of global PAF (8-12 Hz), except for the expected stronger effect of nicotine on change in PAF (𝑏𝑠 = 0.11–0.14, ps < .003; Supplementary Materials 3.8–3.10).

Added in the revised version: Lastly, use of in-house code developed to automatically extract a sensorimotor alpha component using independent component analysis (ICA) showed similar results to the sensorimotor ROI mediation analysis (Supplementary Material 3.11). This supports the robustness of the PAF—pain effect regardless of the method used to isolate sensorimotor alpha.

Summary of mediation analyses

In summary, nicotine demonstrated significant total and direct effects on pain intensity ratings (i.e., reduction compared to placebo) during PHP, but not during CPA, where pain ratings decreased irrespective of gum group. There was also a significant effect of nicotine on change in global PAF speed (i.e., increased compared to placebo). However, there was no evidence of an indirect effect, suggesting that the change in PAF speed did not mediate the change in pain intensity ratings during PHP.

Relationship between PAF and pain at baseline

Earlier version: Evidence for negative correlations between pre-gum PAF (8–12 Hz) and pre-gum pain ratings during PHP (r = −0.14, BF10 = 0.52) and during CPA (r = −0.086, BF10 = 0.33; Figure 5) is inconclusive. Therefore, there was no evidence for or against slower PAF being associated with higher pain ratings. However, in an exploratory analysis (not pre-registered) of separate correlations in males and females (Figure 5, plot C) akin to those conduced in previous research on this topic [10], there was moderate evidence for a negative relationship between global PAF (8–12 Hz) and PHP ratings in males (r = −0.39, BF10 = 3.22). In contrast, the evidence for a relationship between global PAF and PHP ratings in females was equivocal (r = −0.032, BF10 = 0.40). Separating by sex for the correlations between PAF and CPA ratings still produced inconclusive evidence for a relationship (males: r = −0.23, BF10 = 0.61; females: r = −0.086, BF10 = 0.40).

Revised version: Evidence for negative correlations between pre-gum PAF (8–12 Hz) and pre-gum pain ratings during PHP (𝜏 = −0.14, BF10 = 0.52) and during CPA (𝜏 = −0.086, BF10 = 0.33; Figure 5) is inconclusive. Therefore, there was no evidence for or against slower PAF being associated with higher pain ratings. However, in an exploratory analysis (not pre-registered) of separate correlations in males and females (Figure 5, plot C) akin to those conduced in previous research on this topic [10], there was moderate evidence for a negative relationship between global PAF (8–12 Hz) and PHP ratings in males (𝜏 = −0.39, BF10 = 3.22). In contrast, the evidence for a relationship between global PAF and PHP ratings in females was equivocal (𝜏 = −0.032, BF10

Wide band (8–12 Hz) global peak alpha frequency (PAF) does not correlate with A) mean pain ratings during phasic heat pain (PHP) or B) cuff pressure algometry (CPA). Regression lines and shaded 95% confidence intervals. Correlations separated by sex for females in teal (n = 32) and males in brown (n = 30) for C) PHP and D) CPA, suggest a strengthening of a negative relationship between sensorimotor PAF and PHP ratings for males (BF10 = 3.22).

Added in the revised version: = 0.40). Separating by sex for the correlations between PAF and CPA ratings still produced inconclusive evidence for a relationship (males: 𝜏 = −0.23, BF10 = 0.61; females: 𝜏 = −0.086, BF10 = 0.40).

The relationship between PAF and pain was also assessed using median splits on PAF, showing similar patterns of slower PAF being associated with higher PHP ratings in males only (Supplementary Material 6.1). Correlations between global PAF (8–12 Hz) and max pain during PHP also suggested slower PAF was only associated with higher max PHP ratings in male participants (Supplementary Material 6.2). Five participants could not tolerate the 46 °C stimulus during familiarisation, therefore a lower temperature of 45 °C was used for these participants. Repeat analyses excluding these five participants did not impact the observed relationships (Supplementary Material 6.3).

In summary, the evidence for a relationship between PAF and pain is inconclusive, with the exception of the relationship between PAF and PHP ratings by sex, where there was moderate evidence of a negative correlation in male participants.

Discussion

Earlier version: In this pre-registered study, we probed the proposed PAF–pain relationship [8–12] by administration of nicotine – a substance suggested to acutely alter both PAF [46] and pain [51]. Compared to placebo, our data demonstrate three key findings: 1) nicotine increased PAF speed; 2) nicotine reduced prolonged heat pain intensity but not prolonged pressure pain intensity compared to placebo gum; and 3) decreases in prolonged heat pain intensity were not mediated by changes in PAF. This suggests that changes in heat pain due to nicotine occur through mechanisms unrelated to change in PAF. In addition, we found moderate evidence that 4) slower PAF speed at baseline is associated with higher heat pain ratings at baseline in males but not females. In contrast PAF was not associated with pressure pain ratings for males or females.

Revised version: In this pre-registered study, we probed the proposed PAF–pain relationship [8–12] by administration of nicotine – a substance suggested to acutely alter PAF [46]. Compared to placebo, our data demonstrate three key findings: 1) nicotine increased PAF speed; 2) nicotine reduced prolonged heat pain intensity but only when controlling for confounding and not for prolonged pressure pain model; and 3) decreases in pro- longed heat pain intensity were not mediated by changes in PAF. This suggests that changes in heat pain due to nicotine occur through mechanisms unrelated to change in PAF and require replication as the effects required control for confounding and the reduction was small in magnitude. In addition, when assessing the whole sample, we found inconclusive evidence for 4) slower PAF speed at baseline being associated with higher pain ratings for both the heat and pressure pain models. However, during exploratory analysis we found moderate evidence that slower PAF speed at baseline is associated with higher heat pain ratings at baseline in males but not females. In contrast PAF was not associated with pressure pain ratings for males or females.

The effect of nicotine on PAF

Earlier version: We aimed to determine whether nicotine gum could increase PAF speed. In line with previous literature in non-smokers [38,47,48], we contribute additional evidence that chewing 4 mg of nicotine gum increases PAF speed in non-smoking nicotine gum naïve, pain-free, adult participants. A cluster-based permutation analysis indicated that the change in wide window PAF (8–12 Hz) produced by nicotine, was most pronounced in a cluster of electrodes in the central-parietal region and a separate cluster of electrodes in the right-frontal region, rather than in the pre-registered sensorimotor ROI (i.e., Cz, C1, C2, C3, C4).

Revised version: We aimed to determine whether nicotine gum could increase PAF speed. In line with previous literature in non-smokers [38,47,48], we contribute additional evidence that chewing 4 mg of nicotine gum produces a statistically significant increase in PAF speed in non-smoking nicotine gum naïve, pain-free, adult participants. An exploratory cluster-based permutation analysis indicated that the change in wide window PAF (8–12 Hz) produced by nicotine, was most pronounced in a cluster of electrodes in the central-parietal region and a separate cluster of electrodes in the right-frontal region, rather than in the pre-registered sensorimotor ROI (i.e., Cz, C1, C2, C3, C4).

Added in the revised version: Research on nicotine-induced changes in PAF in healthy non-smokers has largely relied on small samples (<20) using injections or patches [47,48], with Bowers and colleagues [38] providing the only larger study (n = 62 males) using 6 mg nicotine gum. Their results showed modest PAF increases (∼0.34–0.51 Hz; 3.5–5.8%) compared to the larger effects (1–2 Hz; ∼10–20%) reported in earlier small-sample studies, with magnitude and scalp distribution influenced by COMT genotype: Met allele carriers exhibited nicotine-induced increases (Val/Met: frontal, central, occipital; Met/Met: parietal), while Val/Val homozygotes showed no change. The present study −— now the largest to include both male and female non-smokers —-found smaller effects, likely due to the lower nicotine dose (4 mg), consistent with dose-dependent PAF modulation reported in abstinent smokers [69]. Broader EEG coverage (64 vs. 8 electrodes), sex inclusion, and unmeasured genetic variation may also account for differences between the present study and Bowers and colleagues [38]. Interestingly, beyond nicotine’s main effect on PAF, our 2W-LCS mediation model indicated a significant influence of sex, with females showing greater PAF increases than males. Given that female participants were significantly lighter than males (Supplementary Material 3.4) and the model did not test the interaction between sex at birth and gum group, these findings should be interpreted cautiously and warrant further investigation.

Removed in the revised version: Interestingly, in addition to effects of nicotine on PAF, the 2W-LCS mediation model suggested that sex also significantly influenced change in PAF, with greater increases in PAF for females compared to males. Notably, the female participants in this study were significantly lighter than the males (Supplementary Material 3.4). However, since the 2W-LCS model did not examine the potential interaction between sex and gum group on change in PAF, conclusions are limited at this stage and future research should investigate these interactions.

The effect of nicotine on pain

Earlier version: Short-term nicotine use is thought to have acute analgesic properties in experimental settings, with a review reporting that nicotine increased pain thresholds and pain tolerance [51]. In addition, research in a rat model suggests analgesic effects on mechanical thresholds after short-term nicotine use [69]. However, previous research has not assessed the effects of nicotine on prolonged experimental pain models. The present study found that 4 mg of nicotine reduced heat pain ratings during prolonged heat pain compared to placebo for our human participants, but that prolonged pressure pain decreased irrespective of which gum was chewed. Our findings are thus partly consistent with the idea that nicotine may have acute analgesic properties [51], although further research is required to explore factors that may influence nicotine’s potential impact on a variety of prolonged pain models. We further advance the literature by reporting this effect in a model of prolonged heat pain, which better approximates the experience of clinical pain than short lasting models used to assess thresholds and tolerance [51].

Revised version: The choice of nicotine was driven by its potential to change PAF, but also, short-term nicotine use is thought to have acute analgesic properties in experimental settings, with a review reporting that nicotine increased pain thresholds and pain tolerance [51]. In addition, research in a rat model suggests analgesic effects on mechanical thresholds after short-term nicotine use [70]. However, previous research has not assessed the effects of nicotine on prolonged experimental pain models. The present study found that 4 mg of nicotine reduced heat pain ratings during prolonged heat pain compared to placebo for our human participants, but that prolonged pressure pain decreased irrespective of which gum was chewed. Our findings are thus partly consistent with the idea that nicotine may have acute analgesic properties [51], although further research is required to explore factors that may influence nicotine’s potential impact on a variety of prolonged pain models. We further advance the literature by reporting this effect in a model of prolonged heat pain, which better approximates the experience of clinical pain than short lasting models used to assess thresholds and tolerance [51].

Earlier version: However, we note that the observed effect of nicotine on pain was small in magnitude, and most prominent in comparison to the effect of placebo, where pain ratings increased after chewing, which brings into question whether this reduction in pain is meaningful in practice. While acknowledging the modest effect size, it’s essential to consider the broader context of our study’s focus. Assessing the clinical relevance of pain reduction is pertinent in applications involving the use of any intervention for pain management [70]. However, from a mechanistic standpoint, particularly in understanding the implications of and relation to PAF, the specific magnitude of the pain effect becomes less pivotal. Nevertheless, future research should examine whether effects on pain increase in magnitude with different nicotine administration regimens (i.e. dose and frequency).

Revised version: However, we note that the observed effect of nicotine on pain was small in magnitude, and most prominent in comparison to the effect of placebo, where pain ratings increased after chewing, which brings into question whether this reduction in pain is meaningful in practice. While acknowledging the modest effect size, it’s essential to consider the context of our study’s focus. Assessing the clinical relevance of pain reduction is pertinent in applications involving the use of any intervention for pain management [71]. However, from a mechanistic standpoint, particularly in understanding the implications of and relation to PAF, the specific magnitude of the pain effect becomes less pivotal, as small effects can still have mechanistic or biological meaning without establishing clinical meaning [72]. Nevertheless, future research should examine whether effects on pain increase in magnitude with different nicotine administration regimens (i.e. dose and frequency).

Earlier version: We did not see an analgesic effect for prolonged pressure pain, as ratings during pressure pain were lower at the second assessment irrespective of gum allocation. The way we calibrated the pressure pain model may have impacted our results. Previous literature has used cuff inflation where the cuff pressure increases until the participant reports 4/10 pain, at which point it remains stable for the desired number of minutes [59–61,71]. In contrast, a separate calibration was conducted at the beginning of this experiment to calculate the pressure used for the remainder of the study, based on the average of three trials of increasing pressure to 4/10 pain (see Methods). This method may result in lower pain ratings at the second pain induction with the same pressure in a same-session repeated measures investigation.

Revised version: We did not see an analgesic effect for prolonged pressure pain, as ratings during pressure pain were lower at the second assessment irrespective of gum allocation. The way we calibrated the pressure pain model may have impacted our results. Previous literature has used cuff inflation where the cuff pressure increases until the participant reports 4/10 pain, at which point it remains stable for the desired number of minutes [59–61,73]. In contrast, a separate calibration was conducted at the beginning of this experiment to calculate the pressure used for the remainder of the study, based on the average of three trials of increasing pressure to 4/10 pain (see Methods). This method may result in lower pain ratings at the second pain induction with the same pressure in a same-session repeated measures investigation.

Earlier version: Analgesic mechanisms of nicotine are complex [51], likely depending on the type of pain and whether nicotine use is acute or chronic, along with several other contributing factors. For example, in addition to the effect of nicotine on prolonged heat pain ratings, our results suggest an effect of stress on changes in heat pain ratings, with those self-reporting higher stress at baseline having greater reductions in pain. Our post-hoc analysis suggested that this relationship between higher stress and larger decrease in PHP ratings was only present for the nicotine group (Supplementary Material 3.3). As stress is linked to nicotine use [72,73] and pain [74–76], these interactions should be explored in future.

Revised version: The mechanisms by which nicotine may influence pain likely depend on the type of pain and whether nicotine use is acute or chronic, along with several other contributing factors [51]. For example, in addition to the effect of nicotine on prolonged heat pain ratings, our results suggest an effect of stress on changes in heat pain ratings, with those self-reporting higher stress at baseline having greater reductions in pain. Our post-hoc analysis suggested that this relationship between higher stress and larger decrease in PHP ratings was only present for the nicotine group (Supplementary Material 3.3). As stress is linked to nicotine use [74,75] and pain [76–78], these interactions should be explored in future.

The relationship between PAF and pain

Earlier version: Prior research suggests that slower PAF is associated with higher pain sensitivity [8–11], and while the validation of PAF as a predictive biomarker continues [14], the question of whether PAF is mechanistically linked to pain remains unanswered. Modulating PAF speed and pain sensitivity with nicotine is necessary but not sufficient to demonstrate a causal pathway from nicotine to PAF to pain (i.e. intervention → mediator → outcome) [77]. We did not find evidence for a causal relationship between PAF and pain sensitivity, as our 2W-LCS model did not show a significant mediating (i.e. indirect) effect of change in PAF on change in pain due to nicotine. The implication of this is that pain sensitivity may not be dependent on state changes in PAF, thus changing PAF may not be a useful target for pain modulation. However, as this is the first study to directly manipulate PAF and observe effects on pain, we cannot confidently state that a causal relationship between PAF and pain sensitivity does not exist, and future research should explore these relationships with larger sample sizes, larger doses of nicotine, and alternative PAF modulators.

Revised version: Prior research suggests that slower PAF is associated with higher pain sensitivity [8–11], and while the validation of PAF as a predictive biomarker continues [13,14], the question of whether PAF is mechanistically linked to pain remains unanswered. Modulating PAF speed and pain sensitivity with nicotine is necessary but not sufficient to demonstrate a causal pathway from nicotine to PAF to pain (i.e. intervention → mediator → outcome) [79]. We did not find evidence for a causal relationship between PAF and pain sensitivity, as our 2W-LCS model did not show a significant mediating (i.e. indirect) effect of change in PAF on change in pain due to nicotine. The implication of this is that pain sensitivity may not be dependent on state changes in PAF, thus changing PAF may not be a useful target for pain modulation. However, as this is the first study attempting to directly manipulate PAF, observing only small effects, we cannot confidently state that a causal relationship between PAF and pain sensitivity does not exist, and future research should explore these relationships with larger sample sizes, larger doses of nicotine, and alternative PAF modulators.

Earlier version: As an alternative explanation, an indirect effect may exist, but could not be seen in the present study due to the single, low dose of nicotine or for other methodological reasons. For example, the relationship between change in PAF and change in pain may be non-linear, and therefore difficult to see using a 2W-LCS model that is based on linear regression analyses. Additionally, the use of global PAF may have introduced mediation measurement error into our mediation analysis. The spatial precision used in the current study was based on previous literature on PAF as a biomarker of pain sensitivity, which have used global and/or sensorimotor ROIs [8,10]. Identification and use of the exploratory electrode clusters found in this study could build upon the current work (e.g., [78]). However, exploratory analysis of the clusters found in the present analysis demonstrated no influence on mediation analysis results (Supplementary Materials 3.8–3.10). Alternatively, independent component analysis (ICA) could be used to identify separate sources of alpha oscillations [79], as used in other experimental PAF–pain studies [8,18], which could aid to disentangle the potential relevance of different alpha sources in the PAF–pain relationship. Although this comes with the need to develop more reproducible and automated methods for identifying such components. In-house code was developed to compare a sensorimotor component to the results presented in this manuscript (Supplementary Material 3.11), showing similar results to the sensorimotor ROI mediation analysis presented here. However, examination of which alpha – be it sensor or source space – are related to pain, how they can be robustly represented, and how they can be manipulated are ripe avenues for future study.

Revised version: As an alternative explanation, an indirect effect may exist, but could not be seen in the present study due to the single, low dose of nicotine or for other methodological reasons. For example, the relationship between change in PAF and change in pain may be non-linear, and therefore difficult to see using a 2W-LCS model that is based on linear regression analyses. Additionally, the use of global PAF may have introduced mediation measurement error into our mediation analysis. The spatial precision used in the current study was based on previous literature on PAF as a biomarker of pain sensitivity, which have used global and/or sensorimotor ROIs [8,10]. Identification and use of the exploratory electrode clusters found in this study could build upon the current work (e.g., [80]). However, exploratory analysis of the clusters found in the present analysis demonstrated no influence on mediation analysis results (Supplementary Materials 3.8–3.10). Alternatively, ICA could be used to identify separate sources of alpha oscillations [81], as used in other experimental PAF–pain studies [8,18], which could aid to disentangle the potential relevance of different alpha sources in the PAF–pain relationship. Although this comes with the need to develop more reproducible and automated methods for identifying such components.

Added in the revised version: In-house code was developed to compare a sensorimotor component to the results presented in this manuscript (Supplementary Material 3.11), showing similar results to the sensorimotor ROI mediation analysis presented here. However, examination of which alpha – be it sensor or source space – are related to pain, how they can be robustly represented, and how they can be manipulated are ripe avenues for future study.

Earlier version: Furthermore, there may have been other mediators that suppress the mediating effects of PAF within our model [66,80,81]. This is plausible, as nicotine is thought to act on pain through multiple pathways [52,54,82], and more direct PAF modulation could more definitively evaluate whether PAF is mechanistically linked to pain sensitivity. Future interventions to modulate PAF could use non-invasive brain stimulation (NIBS), such as transcranial alternative current stimulation (tACS) or repetitive transcranial magnetic stimulation (rTMS), to influence the firing of neurons, thereby directly influencing the brain mechanisms producing PAF [83–85].

Revised version: Furthermore, there may have been other mediators that suppress the mediating effects of PAF within our model [66,82,83]. This is plausible, as nicotine is thought to act on pain through multiple pathways [52,54,84], and more direct PAF modulation could more definitively evaluate whether PAF is mechanistically linked to pain sensitivity. Future interventions to modulate PAF could use non-invasive brain stimulation (NIBS), such as transcranial alternating current stimulation (tACS) or repetitive transcranial magnetic stimulation (rTMS), to influence neuronal firing and thereby directly target the brain mechanisms theorised to generate PAF [85–87]. This approach has since been investigated following the initial preprint of the present manuscript [88,89]. In addition, other approaches to modulating approaches could be considered, such as exercise [32–34], visual stimulation [35,36], cannabis [90,91], and neurofeedback [92,93], while considering their viability in terms of blinding, tolerability, and scalability.

Earlier version: Regarding the relationship between PAF and pain at baseline, the negative correlation between PAF and pain seen in previous work [7–11,16] was only observed here for male participants during the PHP model for global PAF in an exploratory analysis. Effect sizes for the relationship between PAF and pain ratings from previous work were in the range of r = 0.4–0.6 [8–11]. In the present study we saw a weaker relationship for PHP pain (N = 62, r = -0.14). However, sex-stratified analysis demonstrated that males exhibit comparable effect sizes to previous studies (n = 30, r = -0.39), while females did not show a relationship between PAF and heat pain (n = 32, r = -0.032). Subtle variations in experimental design and stimulus delivery could have influenced the PAF–pain relationship. Crucially, the present study did not calibrate the temperature of the PHP model based on participants’ thresholds for moderate pain. Consequently, the fixed temperature did not produce a moderately painful stimulus for all participants, in contrast to previous studies that calibrated to a moderately painful stimuli [10]. This discrepancy may account for the correlation observed between PAF and heat pain ratings among only the male participants in our study, as fixed stimuli can create ceiling or floor effects [86]. Previous literature has also acknowledged the potential for sex differences in the PAF–pain relationship [17].

Revised version: Regarding the relationship between PAF and pain at baseline, the negative correlation between PAF and pain seen in previous work [7–11,16] was only observed here for male participants during the PHP model for global PAF in an exploratory analysis. Effect sizes for the relationship between PAF and pain ratings from previous work were in the range of r = 0.4–0.6 [8–11]. In the present study we saw a weaker relationship for PHP pain (N = 62, r = −0.14). However, sex-stratified analysis demonstrated that males exhibit comparable effect sizes to previous studies (n = 30, r = −0.39), while females did not show a relationship between PAF and heat pain (n = 32, r = −0.032). Subtle variations in experimental design and stimulus delivery could have influenced the PAF–pain relationship. Crucially, the present study did not calibrate the temperature of the PHP model based on participants’ thresholds for moderate pain. Consequently, the fixed temperature did not produce a moderately painful stimulus for all participants, in contrast to previous studies that calibrated to a moderately painful stimuli [10]. This discrepancy may account for the correlation observed between PAF and heat pain ratings among only the male participants in our study, as fixed stimuli can create ceiling or floor effects [94]. Previous literature has also acknowledged the potential for sex differences in the PAF–pain relationship [17].

Earlier version: In parallel, we saw no indication of a relationship between PAF and pain ratings during CPA. The introduction of the CPA model, specifically calibrated to a moderate pain threshold, provides further support for the notion that the relationship between PAF and pain is specific to certain pain types [17,28]. Prolonged heat pain was pre-dominantly described as moderate/severe shooting, sharp, and hot pain, whereas prolonged pressure pain was predominantly described as mild/moderate throbbing, cramping, and aching in the present study. It is possible that the PAF–pain relationship is specific to particular pain models and protocols [12,17].

Revised version: In parallel, we saw no indication of a relationship between PAF and pain ratings during CPA. The introduction of the CPA model, specifically calibrated to a moderate pain threshold, provides further support for the notion that the relationship between PAF and pain could be specific to certain pain types [17,28] or possibly tissues targeted. Prolonged heat pain was pre-dominantly described as moderate/severe shooting, sharp, and hot pain, whereas prolonged pressure pain was predominantly described as mild/moderate throbbing, cramping, and aching in the present study. It is possible that the PAF–pain relationship is specific to particular pain models and protocols [12,17].

Previous experimental pain literature has suggested that slower PAF is related to higher pain sensitivity for PHP [10], capsaicin heat pain [8,10], and nerve growth factor (NGF) [9,13] models, while faster PAF has been related to higher pain sensitivity in tonic heat pain [7] and NGF [15]. Here we contribute to this literature by showing no relationship between baseline PAF and pain sensitivity during the CPA model. Further investigation of the PAF–pain relationship with a variety of prolonged pain models is warranted [12].

Study limitations and constraints

Earlier version: This study presents opportunities for future research due to certain limitations. First, the dose of nicotine administered was equivalent to one cigarette [87]; however, to enhance the efficacy of this research, exploring higher doses or dosing based on weight could be beneficial. Second, blood or saliva samples were not collected in the present study, therefore it is unclear whether participants absorbed similar amounts of nicotine. Future studies in this area should also be powered to perform sex-disaggregated analyses, as our findings indicate a relationship between baseline PAF and subsequent heat pain intensity only in males, and stronger nicotine modulation of PAF in female participants. The ability to detect changes in PAF can be considerably impacted by the frequency resolution used during Fourier Transformations, an element that is overlooked in recent methodological studies on PAF calculation [16,88]. Changes in PAF within individuals might be obscured or conflated by lower frequency resolutions, which should be considered further in future research.

Revised version: This study presents opportunities for future research due to certain limitations. First, the dose of nicotine administered was equivalent to one cigarette [95]; however, to enhance the efficacy of this research, exploring higher doses or dosing based on weight could be beneficial. Second, blood or saliva samples were not collected in the present study, therefore it is unclear whether participants absorbed similar amounts of nicotine. Future studies in this area should also be powered to perform sex-disaggregated analyses, as our findings indicate a relationship between baseline PAF and subsequent heat pain intensity only in males, and stronger nicotine modulation of PAF in female participants. The ability to detect changes in PAF could be considerably impacted by the frequency resolution used during Fourier Transformations, an element that is overlooked in recent methodological studies on PAF calculation [16,96]. Changes in PAF within individuals might be obscured or conflated by lower frequency resolutions, which should be considered further in future research.

Conclusion

In this study, we employed a rigorous double-blind, randomised, placebo-controlled design, pre-registered the analysis, and introduced a novel difference score 2W-LCS mediation model to examine the mediating role of changes in PAF on pain while controlling for pre-determined confounders. These methodological choices allowed for a comprehensive investigation into whether PAF plays a mechanistic role in pain sensitivity.

Earlier version: We have shown that, compared to placebo, nicotine gum increases global PAF speed and decreases pain ratings during prolonged heat pain. However, the effects of nicotine on heat pain were not mediated by changes in PAF. Therefore, we find no evidence that PAF is causally related to pain. The implication of this is that changing PAF may not be a useful target for pain modulation, because pain sensitivity may not be dependent on state changes in PAF. However, this is the first study to test modulation of PAF as a mechanism for acute analgesia. More randomised experiments of the effects of nicotine that use larger sample sizes, track the pharmacokinetics of nicotine, whilst reducing mediator measurement error are needed to uncover whether these factors account for the lack of mediation effect seen here.

Revised version: We have shown that, compared to placebo, nicotine gum produced statistically significant increases global PAF speed and decreases pain ratings during prolonged heat pain. However, the effects of nicotine on heat pain were not mediated by changes in PAF. Therefore, we find no evidence that PAF is causally related to pain. The implication of this is that changing PAF may not be a useful target for pain modulation, because pain sensitivity may not be dependent on state changes in PAF. However, this is the first study to test modulation of PAF as a mechanism for reducing pain intensity. Moreover, while effects on PAF appear small but robust, effects on heat pain required control for confounding to be observed. More randomised experiments on the effects of nicotine that use larger sample sizes, higher nicotine doses, and pharmacokinetic tracking, whilst also reducing mediator measurement error are needed to confirm the present findings and uncover whether these factors account for the lack of mediation effect seen here.

Earlier version: A growing body of evidence indicates a relationship between PAF and pain, including chronic pain studies [17,24,26–28,89], and experimental pain research from our own [8–12,14] and other groups [7,15–18]. Alongside validation of its clinical predictive utility [14,88], further research is needed to ascertain the causal nature of the relationship, explore PAF as a potential target for pain modulation, and gain deeper mechanistic insights into the PAF–pain relationship.

Revised version: A growing body of evidence indicates a relationship between PAF and pain, including chronic pain studies [17,24,26–28,97], and experimental pain research from our own [8–12,14] and other groups [7,15–18]. Alongside validation of its clinical predictive utility [14,96], further research is needed to ascertain the causal nature of the relationship, explore PAF as a potential target for pain modulation, and gain deeper mechanistic insights into the PAF–pain relationship.

Methods

Design overview

The study employed a randomised, double-blind, placebo-controlled, parallel experimental design. Testing sessions were balanced for treatment (i.e. nicotine or placebo gum) and stratified by sex. Individuals were required to attend one 150 minute session. The experimental flow for this study is summarised in Figure 6.

Added in the revised version: Timeline of experimental events.

Earlier version: Timeline of experimental events. Electroencephalography (EEG) data were collected during eyes open and eyes closed resting states before and after each pain assessment. Prolonged pain assessments consisted of cuff pressure algometry (CPA) followed by phasic heat pain (PHP) models, and took place before and after gum chewing.

Revised version: Electroencephalography (EEG) data were collected during eyes open and eyes closed resting states before and after each pain assessment. Prolonged pain assessments consisted of cuff pressure algometry (CPA) followed by phasic heat pain (PHP) models, and took place before and after gum chewing.

Participants

Earlier version: From June 2021 to March 2022, 62 individuals participated at Neuroscience Research Australia (NeuRA). Inclusion criteria for the study were for participants to be healthy, pain free, non-smoking, and between 18–45 years of age, inclusive.

Revised version: From June 2021 to March 2022, 62 individuals participated at Neuroscience Research Australia (NeuRA). Inclusion criteria for the study were for participants to be healthy, pain free, non-smoking, and between 18–45 years of age, inclusive. The sample size was informed by resource and feasibility constraints [98] and aligns with or exceeds that of comparable studies in the field [38], covering 60 participants in total, with two extra in case of incomplete data identified during later data analysis.

Non-smokers were defined as individuals who have smoked fewer than 100 cigarettes and/or e-cigarettes in their lifetime and have not smoked over the past 12 months. Exclusion criteria: previous use of nicotine gum; history of neurological, musculoskeletal, cardiovascular, or psychiatric conditions; other major diseases or disorders; current muscle pain; breastfeeding or pregnant; use of opioids, illicit drugs, or excessive alcohol in the past three months; or unable to provide informed consent.

Ethics approval and consent to participate

The Human Research Ethics Committee (HREC) Executive at the University of New South Wales (UNSW) approved this project (HC190466).

Experimental protocol

Earlier version: Participants were screened for eligibility by telephone and were asked to refrain from drinking coffee or caffeinated drinks two hours prior to the start of testing, as caffeine may impact nicotine absorption [90,91].

Revised version: Participants were screened for eligibility by telephone and were asked to refrain from drinking coffee or caffeinated drinks two hours prior to the start of testing, as caffeine may impact nicotine absorption [99,100].

After signing informed consent, cool and warmth detection tests were conducted, followed by PHP familiarisation, then CPA calibration. During these introductions the participants were also familiarised with the VAS pain rating.

Earlier version: EEG capture started with the control condition of three cycles of warm temperatures, after which the remaining experimental flow is illustrated in Figure 6. The first resting state period was recorded, followed by the first pain assessments (i.e. CPA and PHP model), which were immediately followed by the second resting state. EEG recorded during the pain assessments was not analysed for the present article. Then either nicotine or placebo gum was administered with the chew procedure, followed by the third resting state, the second pain assessments, and the fourth resting state. The experiment ended with participants completing the gum chew experience questionnaire and debrief.

Revised version: EEG capture started with the control condition of three cycles of warm temperatures, after which the remaining experimental flow is illustrated in Figure 6. The first resting state period was recorded, followed by the first pain assessments (i.e. CPA and PHP model), which were immediately followed by the second resting state. EEG recorded during the pain assessments was not analysed for the present article. Participants were randomly assigned to one of only two conditions, thus after the second resting state EEG, either nicotine or placebo gum was administered with the chew procedure. Following completion of the chew procedure, the third resting state, the second pain assessments, and the fourth resting state were conducted. The experiment ended with participants completing the gum chew experience questionnaire and debrief.

Data collection procedures

Questionnaires

Earlier version: Participants completed a set of questionnaires to assess demographics, medical history, general health, mental health, pain, and sleep. This data was collected electronically and stored using REDcap [92,93].

Revised version: Participants completed a set of questionnaires to assess demographics, medical history, general health, mental health, pain, and sleep. This data was collected electronically and stored using REDcap [101,102].

Earlier version: Questionnaires included the following National Institutes of Health common data elements (CDE) for pain biomarkers: Pain Catastrophizing Scale (PCS) [94]; Brief Pain Inventory (BPI), Pain Severity and 7-item Interference subscales [95,96]; SF-8 to assess general health [97]; Patient Health Questionnaire-2 (PHQ-2) to assess depression [98]; Generalised Anxiety Disorder Questionnaire-2 (GAD-2) to assess anxiety [99]; Pennebaker Inventory of Limbic Languidness (PILL) to assess reporting of physical symptoms and sensations [100]; Tobacco, Alcohol, Prescription medications, and other Substances (TAPS); sleep scale. The sleep scale asked “how would you rate your sleep last night?”, on a scale from 1 (i.e. “poor sleep quality”) to 10 (i.e. “excellent sleep quality”). In addition, a health history form, the Karolinska sleepiness scale (KSS) [101], Epworth sleepiness scale (ESS) [102], the perceived stress scale (PSS) [103], and questions regarding participants current level of weekly physical activity were included.

Revised version: Questionnaires included the following National Institutes of Health common data elements (CDE) for pain biomarkers: Pain Catastrophizing Scale (PCS) [103]; Brief Pain Inventory (BPI), Pain Severity and 7-item Interference subscales [104,105]; SF-8 to assess general health [106]; Patient Health Questionnaire-2 (PHQ-2) to assess depression [107]; Generalised Anxiety Disorder Questionnaire-2 (GAD-2) to assess anxiety [108]; Pennebaker Inventory of Limbic Languidness (PILL) to assess reporting of physical symptoms and sensations [109]; Tobacco, Alcohol, Prescription medications, and other Substances (TAPS); sleep scale. The sleep scale asked “how would you rate your sleep last night?”, on a scale from 1 (i.e. “poor sleep quality”) to 10 (i.e. “excellent sleep quality”). In addition, a health history form, the Karolinska sleepiness scale (KSS) [110], Epworth sleepiness scale (ESS) [111], the perceived stress scale (PSS) [112], and questions regarding participants current level of weekly physical activity were included.

Electroencephalography

EEG was recorded from 63 active wet electrodes (actiCap, Brain Products GmbH, Germany), embedded in an elastic cap (EASYCAP, GmbH) in accordance with the international 10-10 system. Recordings were referenced online to ‘FCz’ and the ground electrode was placed on ‘FPz’. A sampling rate of 5000 Hz was used and impedance was kept below 25 kOhms throughout testing using electrolyte gel. EEG signal was amplified and digitised using an actiCHamp Plus DC amplifier linked to BrainVision Recorder software (version 1.22.0101, Brain Products GmbH, Gilching, Germany).

For each recording, resting state EEG was obtained during eyes open (three minutes) followed by eyes closed (three minutes). The eyes open recording took place in low lighting while participants rested their gaze on a white fixation cross with a black background displayed on a monitor. The eyes closed recording took place with the lights off. For both recordings participants were asked to rest in a comfortable seated position and try not to fall asleep.

Pain assessment

Earlier version: Two tonic (i.e. prolonged) pain assessments were conducted during the study, a phasic heat pain (PHP) model [10] and a cuff pressure algometry (CPA) model [59,104]. Pain intensity was sampled at 1 Hz during pain stimulation using a visual analogue scale (VAS), with anchors ‘no pain’ on the left to ‘worst pain imaginable’ on the right, custom MATLAB (version R2021a) code (source available: https://osf.io/pc4rq/).

Revised version: Two tonic (i.e. prolonged) pain assessments were conducted during the study, a phasic heat pain (PHP) model [10] and a cuff pressure algometry (CPA) model [59,113]. Pain intensity was sampled at 1 Hz during pain stimulation using a visual analogue scale (VAS), with anchors ‘no pain’ on the left to ‘worst pain imaginable’ on the right, custom MATLAB (version R2021a) code (source available: https://osf.io/pc4rq/). Participants were instructed to press the ‘STOP’ button on the screen if they became too uncomfortable and wanted to stop the pain assessments.

Removed in the revised version: Participants were instructed to press the ‘STOP’ button on the screen if they became too uncomfortable and wanted to stop the pain assessments.

Earlier version: The Short Form McGill Pain Questionnaire 2 (SF-MPQ-2) [105] was used after each pain model to assess experienced qualities (e.g., throbbing, cramping, burning).

Revised version: The Short Form McGill Pain Questionnaire 2 (SF-MPQ-2) [114] was used after each pain model to assess experienced qualities (e.g., throbbing, cramping, burning).

Heat pain

Added in the revised version: Equipment

Earlier version: Equipment. Thermal stimuli were applied to the volar surface of the left forearm using a thermal-contact heat stimulator (27 mm diameter Medoc Pathways CHEPS Peltier device, Medoc Advanced Medical Systems Ltd., Ramat Yishai, Israel).

Revised version: Thermal stimuli were applied to the volar surface of the left forearm using a thermal-contact heat stimulator (27 mm diameter Medoc Pathways CHEPS Peltier device, Medoc Advanced Medical Systems Ltd., Ramat Yishai, Israel).

Added in the revised version: Cool and warmth detection thresholds

Earlier version: Cool and warmth detection thresholds. The thermode was placed on the volar surface of the left arm at a neutral temperature (32 °C). Participants were instructed to press a button held in their right hand when they sensed a change in temperature. Four areas were marked out by drawing three perpendicular lines equidistant from the elbow crease to 4 cm below the base of the palm. On each trial, the researcher moved through these locations on the forearm in a randomly generated sequence. There were three trials each for warmth (ascending temperatures) and cool (descending temperatures) detection thresholds. Participants were required to consistently detect the warmth and cool transitions in order to continue with the experiment.

Revised version: The thermode was placed on the volar surface of the left arm at a neutral temperature (32 °C). Participants were instructed to press a button held in their right hand when they sensed a change in temperature. Four areas were marked out by drawing three perpendicular lines equidistant from the elbow crease to 4 cm below the base of the palm. On each trial, the researcher moved through these locations on the forearm in a randomly generated sequence. There were three trials each for warmth (ascending temperatures) and cool (descending temperatures) detection thresholds. Participants were required to consistently detect the warmth and cool transitions in order to continue with the experiment.

Added in the revised version: Heat pain threshold

Earlier version: Heat pain threshold. Participants were required to press the hand-held button with their right hand the moment they felt any pain as the thermode temperature increased at a rate of 1.5 °C/second. This pain threshold test was conducted three times and an average heat pain threshold was calculated.

Revised version: Participants were required to press the hand-held button with their right hand the moment they felt any pain as the thermode temperature increased at a rate of 1.5 °C/second. This pain threshold test was conducted three times and an average heat pain threshold was calculated.

Added in the revised version: Familiarisation

Earlier version: Familiarisation. Participants were initially introduced to the PHP model (see below) by receiving two cycles of the heat pain, to ensure they would tolerate the maximum temperature of 46 °C during the experiment and to familiarise participants with using the VAS. If a participant was uncomfortable with receiving 46 °C, the temperature was decreased by 1 °C and the PHP familiarisation was repeated to ensure this temperature was tolerable.

Revised version: Participants were initially introduced to the PHP model (see below) by receiving two cycles of the heat pain, to ensure they would tolerate the maximum temperature of 46 °C during the experiment and to familiarise participants with using the VAS. If a participant was uncomfortable with receiving 46 °C, the temperature was decreased by 1 °C and the PHP familiarisation was repeated to ensure this temperature was tolerable.

Added in the revised version: Phasic warmth

Earlier version: Phasic warmth. Participants experienced three rounds of a phasic warm stimulus. Stimulus trains consisted of a rise phase of 2 °C/s, followed by a 40 second phase where a temperature was held at 36 °C, then a fall phase at a rate of 1 °C/s to reach 32 °C. This neutral temperature was then held for 20 seconds before the cycle repeated three times.

Revised version: Participants experienced three rounds of a phasic warm stimulus. Stimulus trains consisted of a rise phase of 2 °C/s, followed by a 40 second phase where a temperature was held at 36 °C, then a fall phase at a rate of 1 °C/s to reach 32 °C. This neutral temperature was then held for 20 seconds before the cycle repeated three times.

Added in the revised version: Phasic heat pain (PHP) model

Earlier version: Phasic heat pain (PHP) model. The PHP model consisted of a series of five repetitions of phasic stimuli, each lasting one minute 21 seconds. Stimulus trains consisted of a rise phase of 2 °C/s, followed by a 40 second phase where a temperature was held at 46°C, then a fall phase at a rate of 1 °C/s to reach 32 °C. This neutral temperature was then held for 20 seconds before the cycle repeated five times. The PHP model lasted six minutes 45 seconds, and participants rated their pain continuously on the VAS. Similar methodology using repetitive-phasic models have shown that this procedure reliably evokes tonic pain and sensitisation in healthy participants [8,10].

Revised version: The PHP model consisted of a series of five repetitions of phasic stimuli, each lasting one minute 21 seconds. Stimulus trains consisted of a rise phase of 2 °C/s, followed by a 40 second phase where a temperature was held at 46 °C, then a fall phase at a rate of 1 °C/s to reach 32 °C. This neutral temperature was then held for 20 seconds before the cycle repeated five times. The PHP model lasted six minutes 45 seconds, and participants rated their pain continuously on the VAS. Similar methodology using repetitive-phasic models have shown that this procedure reliably evokes tonic pain and sensitisation in healthy participants [8,10].

Pressure pain

Added in the revised version: Equipment

Earlier version: Equipment. Pressure was applied using a standard double-chamber adult arm cuff (length 45cm, width 15cm) placed in the middle of the left lower leg over the head of the gastrocnemius soleus muscle.

Revised version: Pressure was applied using a standard double-chamber adult arm cuff (length 45cm, width 15cm) placed in the middle of the left lower leg over the head of the gastrocnemius soleus muscle.

Added in the revised version: Familiarisation and calibration

Earlier version: Familiarisation and calibration. Participants underwent a pressure cuff calibration. Similar to previous literature [60,62], the pressure was gradually increased manually (30 mmHg/5 seconds) until the participant reported a moderate pain intensity of 4/10 on the VAS and pressure was automatically recorded. This gradual pressure increase was conducted three times with a two minute break between each test, to establish the average pressure required to evoke 4/10 pain for each participant. The maximum pressure limit for this study was 100 kPa (750 mmHg) [63]. This is well within safe limits, as previous studies have used maximums of 180 or 200 kPa without adverse effects [59–61].

Revised version: Participants underwent a pressure cuff calibration. Similar to previous literature [60,62], the pressure was gradually increased manually (30 mmHg/5 seconds) until the participant reported a moderate pain intensity of 4/10 on the VAS and pressure was automatically recorded. This gradual pressure increase was conducted three times with a two minute break between each test, to establish the average pressure required to evoke 4/10 pain for each participant. The maximum pressure limit for this study was 100 kPa (750 mmHg) [63]. This is well within safe limits, as previous studies have used maximums of 180 or 200 kPa without adverse effects [59–61].

Added in the revised version: Cuff pressure algometry (CPA) model

Earlier version: Cuff pressure algometry (CPA) model. Pressure was increased manually (30 mmHg/5 seconds) until the individualised pressure was reached, then held constant for four minutes. Participants rated their pain continuously using the VAS.

Revised version: Pressure was increased manually (30 mmHg/5 seconds) until the individualised pressure was reached, then held constant for four minutes. Participants rated their pain continuously using the VAS.

Gum chewing

Earlier version: Nicotine was administered orally in the form of 4 mg of cinnamon flavoured Nicorette® gum (Johnson & Johnson Inc., Markham, Ontario, Canada); comparable to smoking a single cigarette, this produces a blood concentration of approximately 15-30 mg/ml [87]. The placebo gum was cinnamon flavoured Xyloburst® natural chewing gum, which, other than containing no nicotine, is of a similar flavour, size, and texture to the nicotine gum. At the end of the experiment, before being debriefed, participants completed a short questionnaire about the gum chewing experience to test perceptual matching between the nicotine and placebo gums.

Revised version: Nicotine was administered orally in the form of 4 mg of cinnamon flavoured Nicorette® gum (Johnson & Johnson Inc., Markham, Ontario, Canada); comparable to smoking a single cigarette, this produces a blood concentration of approximately 15-30 mg/ml [95]. The placebo gum was cinnamon flavoured Xyloburst® natural chewing gum, which, other than containing no nicotine, is of a similar flavour, size, and texture to the nicotine gum. At the end of the experiment, before being debriefed, participants completed a short questionnaire about the gum chewing experience to test perceptual matching between the nicotine and placebo gums.

Earlier version: To promote nicotine absorption, participants rinsed their mouth with a glass of water for 10 seconds prior to receiving the gum, as performed previously [90]. We used a protocol similar to Bowers and colleagues [38], following manufacturer’s guidelines of a total chewing time of 25 minutes. Audio cues to bite the gum occurred every 30 seconds for a total of 25 minutes; participants were asked to alternate between the left and right sides of the mouth, and to park the gum between the teeth and cheek between chews. Participants watched a wildlife documentary during the chewing to maintain attention; all participants watched the same documentary to control for its possible effects.

Revised version: To promote nicotine absorption, participants rinsed their mouth with a glass of water for 10 seconds prior to receiving the gum, as performed previously [99]. We used a protocol similar to Bowers and colleagues [38], following manufacturer’s guidelines of a total chewing time of 25 minutes. Audio cues to bite the gum occurred every 30 seconds for a total of 25 minutes; participants were asked to alternate between the left and right sides of the mouth, and to park the gum between the teeth and cheek between chews. Participants watched a wildlife documentary during the chewing to maintain attention; all participants watched the same documentary to control for its possible effects.

Randomisation and blinding

Participants were randomised in a 1:1 ratio to receive either nicotine gum or placebo gum. The allocation list was created a priori, independently of the primary experimenter, using blocked randomisation based on sex at birth. The allocation was blind to both the participant and experimenter. All participants were told that they were receiving nicotine gum to provide equal expectations. Participants were debriefed at the end of the experiment and asked whether they thought they had been given nicotine, placebo, or if they weren’t sure. Participants were unblinded by the independent researcher separately. The primary experimenter remained blinded with dummy variables until completion of analysis.

Data processing

EEG pre-processing

Earlier version: Blinded preprocessing of de-identified EEG was conducted with custom code in MATLAB 2020a (v.9.8.0.1417392, The MathWorks, Inc., Natick, MA, USA) using EEGLAB 2020.0 [106] and FieldTrip (v.20200607) [107].

Revised version: Blinded preprocessing of de-identified EEG was conducted with custom code in MAT-LAB 2020a (v.9.8.0.1417392, The MathWorks, Inc., Natick, MA, USA) using EEGLAB 2020.0 [115] and FieldTrip (v.20200607) [116].

Earlier version: Using EEGLAB, for four resting states, each participant’s data were downsampled to 500 Hz, filtered between 2 and 100 Hz using a linear finite impulse response (FIR) bandpass filter, and the three minute periods of eyes closed resting state were visually inspected and overtly noisy channels with ≥ 8 signal deviations (e.g., popping, baseline drift, line noise) were removed. The average number of channels removed were 0.94, 0.98, 0.74, and 0.87 (range: 0–4) for the four resting states, respectively. Data from the remaining channels were re-referenced to the common average reference. Removed channels were not interpolated for the pre-registered global and sensorimotor ROI averaged analyses, but were interpolated for an exploratory cluster based permutation analysis using the nearest neighbour average method in Fieldtrip.

Revised version: Using EEGLAB, for four resting states, each participant’s data were downsampled to 500 Hz, filtered between 2 and 100 Hz using a linear finite impulse response (FIR) bandpass filter, and the three minute periods of eyes closed resting state were visually inspected and overtly noisy channels with ≥ 8 signal deviations (e.g., popping, base-line drift, line noise) were removed. The average number of channels removed were 0.94, 0.98, 0.74, and 0.87 (range: 0–4) for the four resting states, respectively. Data from the remaining channels were re-referenced to the common average reference. Removed channels were not interpolated for the pre-registered global and sensorimotor ROI averaged analyses, but were interpolated for an exploratory cluster based permutation analysis using the nearest neighbour average method in Fieldtrip.

Earlier version: Using FieldTrip, data were segmented into 5 second epochs, with no overlap, visually inspected using ft_databrowser(), and epochs containing any electrode, muscle, or motion artefacts were removed. From a maximum of 36 epochs (i.e. 3 mins), the average number of remaining epochs was 31.21, 31.33, 32.13, and 31.39 (range: 24–36) for each resting state, respectively. No resting state contained less than two minutes of data [88]. Principal component analysis using the runica method was applied to identify and remove components reflecting eye blinks and/or saccades. The average number of components removed per resting state was 1.05, 1.2, 1.1, 1.2 (range: 0–2), respectively.

Revised version: Using FieldTrip, data were segmented into 5 second epochs, with no overlap, visually inspected using ft_databrowser(), and epochs containing any electrode, muscle, or motion artefacts were removed. From a maximum of 36 epochs (i.e. 3 mins), the average number of remaining epochs was 31.21, 31.33, 32.13, and 31.39 (range: 24–36) for each resting state, respectively. No resting state contained less than two minutes of data [96]. Principal component analysis using the runica method was applied to identify and remove components reflecting eye blinks and/or saccades. The average number of components removed per resting state was 1.05, 1.2, 1.1, 1.2 (range: 0–2), respectively.

Quantification of peak alpha frequency (PAF)

Earlier version: Each participant’s PAF was estimated for the four eyes closed resting states. FieldTrip was used to calculate the power spectral density in the range of 2–40 Hz for each epoch in 0.2 Hz bins. A Hanning taper was applied to the data prior to calculating spectra to reduce any edge artefacts [108]. The PAF for each power spectral density was estimated using the centre of gravity (CoG) method [109–113]: where fiis the ith frequency bin, n is the number of frequency bins within the frequency window (8–12 Hz or 9–1 Hz), and ai is the amplitude of the frequency for fi. This equation states that PAF is the weighted sum of spectral estimates [109]. Global (i.e. average of all electrodes) and a sensorimotor region of interest were used (i.e. average of C3, C1, Cz, C2, and C4). PAF estimates for each participant at each time point were calculated. Out of 248 resting states recorded, 14 resting states had 4 ROI channels instead of 5. Importantly, no resting state had fewer than 4 channels for the sensorimotor ROI.

Revised version: Each participant’s PAF was estimated for the four eyes closed resting states. FieldTrip was used to calculate the power spectral density in the range of 2–40 Hz for each epoch in 0.2 Hz bins. A Hanning taper was applied to the data prior to calculating spectra to reduce any edge artefacts [117]. The PAF for each power spectral density was estimated using the centre of gravity (CoG) method [118–122]: (1) CoG=∑i=1nfiai∑i=1nai where 𝑓𝑖 is the 𝑖th frequency bin, 𝑛 is the number of frequency bins within the frequency window (8–12 Hz or 9–1 Hz), and 𝑎𝑖 is the amplitude of the frequency for 𝑓𝑖. This equation states that PAF is the weighted sum of spectral estimates [118]. Global (i.e. average of all electrodes) and a sensorimotor region of interest were used (i.e. average of C3, C1, Cz, C2, and C4). PAF estimates for each participant at each time point were calculated. Out of 248 resting states recorded, 14 resting states had 4 ROI channels instead of 5. Importantly, no resting state had fewer than 4 channels for the sensorimotor ROI.

Contrary to some previous studies that used ICA to isolate sensory region alpha sources [8,15,18], we used pre-determined sensor level ROIs to improve reproducibility and reduce the potential for bias when individually selecting ICA components. Using sensor level ROIs may decrease the signal-to-noise ratio of the data; however, this approach has still been effective for observing the relationship between PAF and experimental pain [9,10].

Earlier version: The frequency range used to calculate PAF varies in the PAF–pain literature. The narrower 9–11 Hz band reduces the impact of 1/f noise (i.e. trend of higher power at lower frequencies) on the PAF estimations [8,108], however the narrow band may not produce representative PAF values for those with peaks near the lower or upper limits. Therefore, the wider 8–12 Hz band was used as the primary band of interest (Figure 2, plot A). All analyses were repeated with the narrower band and are displayed in Supplementary Material 3.4 and 3.6. Power was calculated using the wider 8–12 Hz band to display alpha on topographical plots.

Revised version: The frequency range used to calculate PAF varies in the PAF–pain literature. The narrower 9–11 Hz band reduces the impact of 1/f noise (i.e. trend of higher power at lower frequencies) on the PAF estimations [8,117], however the narrow band may not produce representative PAF values for those with peaks near the lower or upper limits. Therefore, the wider 8–12 Hz band was used as the primary band of interest (Figure 2, plot A). All analyses were repeated with the narrower band and are displayed in Supplementary Material 3.4 and 3.6. Power was calculated using the wider 8–12 Hz band to display alpha on topographical plots.

Statistical analysis

Earlier version: Statistical analysis was conducted in RStudio (RStudio Team, 2021.9.1.372). A Guideline for Reporting Mediation Analysis (AGReMA) was followed for this article and the analysis plan was pre-registered (https://osf.io/pc4rq/).

Revised version: Statistical analysis was conducted in RStudio (RStudio Team, 2021.9.1.372). A Guideline for Reporting Mediation Analysis (AGReMA) was followed for this article [123] and the analysis plan was pre-registered (https://osf.io/pc4rq/).

Mediation analysis

Using a structural equation modelling approach, a two-wave latent change score (2W-LCS) model – recommended by Valente and MacKinnon [64] – was used to assess whether nicotine changed PAF or pain (i.e. hypotheses 1–3).

Earlier version: Prior to modelling, a paired equivalence test satisfied that PAF1a (i.e. pre-gum before the first pain assessment) and PAF1b (i.e. pre-gum after the first pain assessment) were equivalent (t(60) = 0.90, p < .001), therefore PAF1a was used in the below models.

Revised version: Prior to modelling, a paired equivalence test satisfied that PAF1𝑎 (i.e. pre-gum before the first pain assessment) and PAF1𝑏 (i.e. pre-gum after the first pain assessment) were equivalent (𝑡(60) = 0.90, 𝑝 < .001), therefore PAF1𝑎 was used in the below models.

Added in the revised version: Latent change score (LCS) specification

Earlier version: Latent change score (LCS) specification. Calculating the absolute difference between pre- and post-gum variables can result in regression to the mean [64]. With LCS specification, the second wave of PAF and pain (i.e. post-gum) can be set as a unit- weighted function of pre-gum and a new unobserved change variable (i.e. ΔPAF and Δpain) to simulate subtracting pre-gum scores from post-gum scores on PAF and pain variables [64].

Revised version: Calculating the absolute difference between pre- and post-gum variables can result in regression to the mean [64]. With LCS specification, the second wave of PAF and pain (i.e. post-gum) can be set as a unitweighted function of pre-gum and a new unobserved change variable (i.e. Δ𝑃𝐴𝜞 and Δ𝑝𝑎𝑖𝑛) to simulate subtracting pre-gum scores from post-gum scores on PAF and pain variables [64]. (2) ΔPAF=PAF2a−PAF1a

Earlier version: Simulation of the above equations was achieved via the following parameter constraints. Loading on Pain2 was fixed to one, while the mean and variance of Δpain were freely estimated. The path from Pain1 to Pain2 was fixed to 1, the mean (i.e. intercept) and variance of Pain2 were constrained to zero, and the mean and variance of Pain1 were freely estimated. The same approach was taken for ΔPAF . The covariances between PAF1a and Pain1 were freely estimated.

Revised version: Simulation of the above equations was achieved via the following parameter constraints. Loading on Pain2 was fixed to one, while the mean and variance of Δ𝑝𝑎𝑖𝑛 were freely estimated. The path from Pain1 to Pain2 was fixed to 1, the mean (i.e. intercept) and variance of Pain2 were constrained to zero, and the mean and variance of Pain1 were freely estimated. The same approach was taken for Δ𝑃𝐴𝜞 . The covariances between PAF1𝑎 and Pain1 were freely estimated.

Added in the revised version: Two-wave latent change score (2W-LCS) models

Earlier version: Two-wave latent change score (2W-LCS) models. The above latent change variables (Equation 2 and Equation 3) can then be used within 2W-LCS models to conduct mediation analysis. Two model variations were constructed to assess total, direct, and indirect effects of nicotine separately for latent change in mean ratings during PHP and CPA as outcomes (Table 4). Models were constructed in R (version 4.2.1) using the lavaan package [116] with maximum likelihood estimation and 1000 bootstraps (seed: 2022) to compute standard errors. The mediation model syntax is available online (https://osf.io/pc4rq/).

Revised version: The above latent change variables (Equation 2 and Equation 3) can then be used within 2W-LCS models to conduct mediation analysis. Two model variations were constructed to assess total, direct, and indirect effects of nicotine separately for latent change in mean ratings during PHP and CPA as outcomes (Table 4). Models were constructed in R (version 4.2.1) [124] using the lavaan package [125] with maximum likelihood estimation and 1000 bootstraps (seed: 2022) to compute standard errors. The mediation model syntax is available online (https://osf.io/pc4rq/).

Added in the revised version: The ANCOVA-equivalent 2W-LCS model

Earlier version: The ANCOVA-equivalent 2W-LCS model. The first model variation used was the ANCOVA-equivalent model [64], the path diagram is displayed in Figure 7 and the below equations were used (Equation 4 and Equation 5):

Revised version: The first model variation used was the ANCOVA-equivalent model [64], the path diagram is displayed in Figure 7 and the below equations were used (Equation 4 and Equation 5): (4) ΔPAF∼i1+am2xGum+sm2m1PAF1+bm2y1 Pain 1+covariates+e1 (5) Δpain∼i2+cy2y1′Gum+sy2y1 Pain 1+by2m1PAF1+by2m2ΔPAF+covariates+e2

Added in the revised version: The analysis of covariance (ANCOVA) equivalent pre-test/post-test control group design (two-wave) with latent change score (LCS) specification for the outcome (Y: pain) and mediating (M: PAF) variables.

Removed in the revised version: In this model, i denotes the intercept and e denotes error. Both sm2m1M1 and sy2y1Y1 are time-lagged autoregressions referring to change in M and Y due to prior levels of M and Y, respectively. These autoregressions reflect that each variable influences itself over time. The by2m1M1 and bm2y1Y1 are cross-lagged coupling parameters, which refer to effects of previous levels of M on change in Y and effects of previous levels of Y on change in M. In other words, each variable influences the other variable at the subsequent time point. As displayed in Figure 7, X is predicting constant change in Y via the direct c’–path (Equation 5; c′) and constant change in M via the a–path (Equation 4; am2x). Latent change in M is predicting later latent change in Y via the constrained b–path (Equation 5; by2m2). The indirect effect was assessed by estimating the product of coefficients ab (am2xX ∗ by2m2). The total effect was assessed by summation of the indirect and direct effects (i.e. c′ + [am2xX ∗ by2m2]).

Removed in the revised version: Adapted from Valente and MacKinnon [64]. The analysis of covariance (ANCOVA) equivalent pre-test/post-test control group design (two-wave) with latent change score (LCS) specification for the outcome (Y: pain) and mediating (M: PAF) variables. The direct effect of the intervention on the outcome is represented by the red line (c′). The intervention–mediator relationship is represented by a green line (am2x) from the intervention (i.e. X: nicotine vs. placebo) to the mediator (i.e. ΔM). The mediator–outcome relationship is represented by another green line (by2m2) from the mediator (i.e. ΔM) to the outcome (i.e ΔY). Together the two green lines represent the indirect effect of the intervention on the outcome via the mediator. Time-lagged autoregression paths are represented by S∗m2m1 and S∗y2y1. Cross-lagged coupling paths are represented by bm2y1 and b∗y2m1. The covariance between Y and M at baseline is represented by σm1y1. The potential confounders of the mediator–outcome relationship were measured at baseline and are represented by the blue dashed lines. Note. ∗ = latent; X = intervention (i.e. gum); M = mediator (i.e. PAF); Y = outcome (i.e. pain).

Removed in the revised version: Difference score 2W-LCS model. We compared the ANCOVA-equivalent model (Figure 7) to the difference score model (Figure 8); the former adjusts for the effects of pre-gum PAF and pain on post-gum PAF and pain, while the latter assumes those paths are equal to zero [64]. Removal of autoregressive (i.e. sm2m1M1 and sy2y1Y1) and coupling parameters (i.e. by2m1M1 and bm2y1Y1) produces a simpler unconditional model of change [65,117]. The path diagram is displayed in Figure 8 and the below equations were used (Equation 6 and Equation 7):

Removed in the revised version: Model fit and comparison. Model fits were compared with χ2 statistic (p >0.05), comparative fit index (CFI), standardised root mean squared residual (SRMR), and root mean square error of approximation (RMSEA) [65,118]. For adequate model fit: CFI > 0.90, RMSEA < 0.08, and SRMR < 0.08 [119].

Earlier version: Adapted from Valante and MacKinnon [64]. The difference score model with latent change score specification for the outcome (Y) and mediating (M) variables. The direct effect of the intervention on the outcome is represented by the red line (c′). The intervention–mediator relationship is represented by a green line (am2x) from the intervention (i.e. X: nicotine vs. placebo) to the mediator (i.e. ΔM). The mediator–outcome relationship is represented by another green line (bΔ) from the mediator (i.e. ΔM) to the outcome (i.e ΔY). Together the two green lines represent the indirect effect of the intervention on the outcome via the mediator. Cross-lagged coupling paths are constrained to 0 and time-lagged autoregressions are constrained to 1. The covariance between Y and M at baseline is represented by σm1y1. The potential confounders of the mediator–outcome relationship are represented by the blue dashed lines. Note. X = intervention (i.e. gum); M = mediator (i.e. PAF); Y = outcome (i.e. pain).

Revised version: The direct effect of the intervention on the outcome is represented by the red line (𝑐’𝑦2𝑥). The intervention–mediator relationship is represented by a green line (𝑎𝑚2𝑥) from the intervention (i.e. X: nicotine vs. placebo) to the mediator (i.e. Δ𝑀). The mediator–outcome relationship is represented by another green line (𝑏𝑦2𝑚2) from the mediator (i.e. Δ𝑀) to the outcome (i.e Δ𝑌). Together the two green lines represent the indirect effect of the intervention on the outcome via the mediator. Time-lagged autoregression paths are represented by 𝑆∗𝑚2𝑚1 and 𝑆∗𝑦2𝑦1. Cross-lagged coupling paths are represented by 𝑏𝑚2𝑦1 and 𝑏∗𝑦2𝑚1. The covariance between Y and M at baseline is represented by 𝜎𝑚1𝑦1. The potential confounders of the mediator–outcome relationship were measured at baseline and are represented by the blue dashed lines. Note. ∗ = latent; X = intervention (i.e. gum); M = mediator (i.e. PAF); Y = outcome (i.e. pain). Adapted from Valente and MacKinnon [64].

Added in the revised version: In this model, 𝑖 denotes the intercept and 𝑒 denotes error. Both 𝑠𝑚2𝑚1𝑀1 and 𝑠𝑦2𝑦1𝑌1 are time-lagged autoregressions referring to change in M and Y due to prior levels of M and Y, respectively. These autoregressions reflect that each variable influences itself over time. The 𝑏𝑦2𝑚1𝑀1 and 𝑏𝑚2𝑦1𝑌1 are cross-lagged coupling parameters, which refer to effects of previous levels of M on change in Y and effects of previous levels of Y on change in M. In other words, each variable influences the other variable at the subsequent time point. As displayed in Figure 7, X is predicting constant change in Y via the direct c’–path (Equation 5; cy2y1′) and constant change in M via the a–path (Equation 4; 𝑎𝑚2𝑥). Latent change in M is predicting later latent change in Y via the constrained b–path (Equation 5; 𝑏𝑦2𝑚2). The indirect effect was assessed by estimating the product of coefficients 𝑎𝑏 (𝑎𝑚2𝑥𝐸 ∗ 𝑏𝑦2𝑚2). The total effect was assessed by summation of the indirect and direct effects cy2y1′+[am2xX∗by2m2]).

Added in the revised version: Difference score 2W-LCS model

Added in the revised version: We compared the ANCOVA-equivalent model (Figure 7) to the difference score model (Figure 8); the former adjusts for the effects of pre-gum PAF and pain on post-gum PAF and pain, while the latter assumes those paths are equal to zero [64]. Removal of autoregressive (i.e. 𝑠𝑚2𝑚1𝑀1 and 𝑠𝑦2𝑦1𝑌1) and coupling parameters (i.e. 𝑏𝑦2𝑚1𝑀1 and 𝑏𝑚2𝑦1𝑌1) produces a simpler unconditional model of change [65,126]. The path diagram is displayed in Figure 8 and the below equations were used (Equation 6 and Equation 7): (6) Δm∼i3+aGum+covariates+e3 (7) Δy∼i4+c′Gum+bΔPAF+covariates+e4

Added in the revised version: The difference score model with latent change score specification for the outcome (Y) and mediating (M) variables.

Added in the revised version: The direct effect of the intervention on the outcome is represented by the red line (𝑐’Δ). The intervention–mediator relationship is represented by a green line (𝑎𝑚2𝑥) from the intervention (i.e. X: nicotine vs. placebo) to the mediator (i.e. Δ𝑀). The mediator–outcome relationship is represented by another green line (𝑏Δ) from the mediator (i.e. Δ𝑀) to the outcome (i.e Δ𝑌). Together the two green lines represent the indirect effect of the intervention on the outcome via the mediator. Cross-lagged coupling paths are constrained to 0 and time-lagged autoregressions are constrained to 1. The covariance between Y and M at baseline is represented by 𝜎𝑚1𝑦1. The potential confounders of the mediator–outcome relationship are represented by the blue dashed lines. Note. X = intervention (i.e. gum); M = mediator (i.e. PAF); Y = outcome (i.e. pain). Adapted from Valante and MacKinnon [64].

Added in the revised version: Model fit and comparison

Added in the revised version: Model fits were compared with 𝜒2statistic (p >0.05), comparative fit index (CFI), standardised root mean squared residual (SRMR), and root mean square error of approximation (RMSEA) [65,127]. For adequate model fit: CFI > 0.90, RMSEA < 0.08, and SRMR < 0.08 [128].

Added in the revised version: Confounding

Earlier version: Confounding. The intervention–mediator (a–path) and intervention–outcome (c’–path) relationships are unconfounded due to the random allocation of nicotine or placebo gum. For the mediator–outcome (b–path) relationship, potential confounders were chosen by selecting variables measured prior to gum allocation that we assume to be a cause of the mediator, outcome, or both [120].

Revised version: The intervention–mediator (a–path) and intervention–outcome (c’– path) relationships are unconfounded due to the random allocation of nicotine or placebo gum. For the mediator–outcome (b–path) relationship, potential confounders were chosen by selecting variables measured prior to gum allocation that we assume to be a cause of the mediator, outcome, or both [129].

Earlier version: We assume that the mediator–outcome relationship is unconfounded after adjusting for the pre-determined (Supplementary Material 7.1; [121,122]) minimum sufficient set of observed confounders: sex at birth, depression symptoms, anxiety symptoms, perceived stress, sleep quality, and recent pain (Figure 7). Within the model, the variance of each covariate was freely estimated, along with their co-variance with baseline values of PAF and pain.

Revised version: We assume that the mediator–outcome relationship is unconfounded after adjusting for the pre-determined (Supplementary Material 7.1; [130,131]) minimum sufficient set of observed confounders: sex at birth, depression symptoms, anxiety symptoms, perceived stress, sleep quality, and recent pain (Figure 7). Within the model, the variance of each covariate was freely estimated, along with their co-variance with baseline values of PAF and pain.

Added in the revised version: Assumptions

Earlier version: Assumptions. Multicollinearity assessment (R2 cut-off of 0.7) did not suggest any covariates requiring removal. Using the mvn() function, we found violation of multivariate normality (Mardia skewness = 361.28, p <.001; Mardia kurtosis = 2.39, p = 0.017) as well as univariate normality for covariates 2, 3, 5, and 6 (W = 0.60-0.93, p < 0.0019), the impact of which was reduced by use of bootstrapping [123,124]. There was no systematically missing data and only complete datasets were used, incomplete sets were removed by listwise deletion.

Revised version: Multicollinearity assessment (𝑅2 cut-off of 0.7) did not suggest any covariates requiring removal. Using the mvn() function, we found violation of multivariate normality (Mardia skewness = 361.28, p <.001; Mardia kurtosis = 2.39, p = 0.017) as well as univariate normality for covariates 2, 3, 5, and 6 (W = 0.60-0.93, p < 0.0019), the impact of which was reduced by use of bootstrapping [132,133]. There was no systematically missing data and only complete datasets were used, incomplete sets were removed by listwise deletion.

Added in the revised version: Mediation using PROCESS

Earlier version: Mediation using PROCESS. PROCESS is a statistics package in R for conducting mediation analysis [66], and is frequently used in the literature [125–127]. PROCESS analysis was conducted using and ANCOVA-equivalent model without the LCS specification, as construction of latent variables is not currently possible in this package [66]. Pain and PAF post-gum (i.e. PAF2a and pain2) were used as outcome measures rather than change in each, with PAF1 and pain1 included alongside the other covariates. The estimates produced by PROCESS were similar to the ANCOVA-equivalent 2W-LCS model in lavaan, and are displayed in the Supplementary Material 7.2.

Revised version: PROCESS is a statistics package in R for conducting mediation analysis [66], and is frequently used in the literature [134–136]. PROCESS analysis was conducted using and ANCOVA-equivalent model without the LCS specification, as construction of latent variables is not currently possible in this package [66]. Pain and PAF post-gum (i.e. PAF2𝑎 and pain2) were used as outcome measures rather than change in each, with PAF1 and pain1 included alongside the other covariates. The estimates produced by PROCESS were similar to the ANCOVA-equivalent 2W-LCS model in lavaan, and are displayed in the Supplementary Material 7.2.

Added in the revised version: Cluster-based permutation analysis

Earlier version: Cluster-based permutation analysis. An exploratory non-parametric cluster-based permutation analysis [128,129] was conducted to assess where changes in global PAF occurred in sensor space, similar to those conduced by Furman and colleagues [78]. Tests conducted at the single electrode level were thresholded at p < .01. If two or more spatially adjacent electrodes passed this threshold, they were grouped into a cluster, and each electrode within a cluster had to be adjacent to at least one other electrode from the overall cluster. Following the formation of clusters, test statistic values were summed across all electrodes within a cluster to obtain a composite test score for that cluster. For each potential cluster, this composite test score was then compared against a null distribution of 1000 generated scores produced by random shuffle of values for the dependent variable across conditions and participants at each electrode of the cluster. If a cluster’s composite test statistic surpassed the 99th percentile of the null distribution, that cluster was evaluated as significant.

Revised version: An exploratory non-parametric cluster-based permutation analysis [137,138] was conducted to assess where changes in global PAF occurred in sensor space, similar to those conduced by Furman and colleagues [80]. Tests conducted at the single electrode level were thresholded at 𝑝 < .01. If two or more spatially adjacent electrodes passed this threshold, they were grouped into a cluster, and each electrode within a cluster had to be adjacent to at least one other electrode from the overall cluster. Following the formation of clusters, test statistic values were summed across all electrodes within a cluster to obtain a composite test score for that cluster. For each potential cluster, this composite test score was then compared against a null distribution of 1000 generated scores produced by random shuffle of values for the dependent variable across conditions and participants at each electrode of the cluster. If a cluster’s composite test statistic surpassed the 99th percentile of the null distribution, that cluster was evaluated as significant.

PAF–pain relationship

To address whether faster sensorimotor or global PAF during pain-free resting states was associated with lower pain ratings during the first PHP and CPA models before gum chewing (i.e. hypothesis 4), correlations were used [7–10].

Bayesian correlations were conducted in R (‘BayesFactor’ package, default priors) between pain-free baseline 8–12 Hz PAF (i.e. global and sensorimotor) and pain measures (i.e. mean, max, AUC) during each model (i.e. PHP, CPA) prior to gum administration.

Median splits on PAF speed were conducted, and the difference in pain reported between PAF speed groups was assessed by independent t-tests for mean and AUC pain, with Mann Whitney-U test for max PHP pain, due to violation of normality (W = 0.89, p < .001). Results of median splits are displayed in Supplementary Material 6.1.